Knockout of cyclin dependent kinases 8 and 19 leads to depletion of cyclin C and suppresses spermatogenesis and male fertility in mice

Author response:

The following is the authors’ response to the original reviews.

The mice crossing scheme is unusual as you have three mice to cross to produce genotypes, while we could understand that it is possible to produce pups of desired genotypes with different mating schemes, such a vague crossing scheme is not desirable and of poor genetics practice.

We thank the reviewer for this suggestion. Indeed, our scheme is not a representation of the actual breeding scheme but just a brief explanation of lineages used for the acquisition of the triple transgenic mice. We will include the full crossing scheme into the revision.

We added to the text the explanation that all used genotypes were maintained as homozygotes and put a full breeding scheme in the supplementary figure S1A

It is worth mentioning that single knockouts seem to show a corresponding upregulation of the level of the paralogue kinase, indicating that any lack of phenotypes might be due to feedback compensation, which would be an interesting finding if confirmed; this has not been mentioned.

We thank the reviewer for raising an important point about the paralog upregulation. Indeed, our data on primary cells (supplementary 1B) suggests the upregulation of CDK19 in CDK8KO and vice versa. We will point this out in discussion. We plan to examine the data for the testis as soon as more tissues are available.

We addressed this question by performing additional western blot (added to the paper fig. 2D) and found no paralogue upregulation in testes. To do that we also manufactured novel rabbit anti-mouse CDK19 antibodies described in Materials and Methods.

The authors should clarify or present the data on where CDK8 and CDK19 as well as CcnC are expressed so as to help the readers understand which tissues both CDK might be functioning in and cause the loss of CcnC.

Due to a limited sensitivity of single cell sequencing (only ~5,000 transcripts are sequenced from total of average 500,000 transcripts per cell, so the low expressed transcripts are not sequenced in all cells) it is challenging to firmly establish CDK8/19 positive and -negative tissues from single cell data because both transcripts are minor. This image will be included in the next version.

In this version we have added staining by CDK8 and CDK19 antibodies on paraffin sections, showing expression in variety of cells. Additionally, we have analyzed Cdk8/CcnC presence in different testicular cell types by flow cytometry. Both methods show that not only spermatogonial stem cells express Cdk8 as was shown in McCleland et al. 2005, but also some 1n cells, 4n cells and a significant part of cKit- 2n cells. We added a corresponding paragraph and figures (2E-K) to the paper. We consider this a more definitive answer to the question than RNA data.

Furthermore, data for the genitourinary system in single knockouts are very sparse; data are described for fertility in Figure 1H, ploidy, and cell number in Figures 2B and C, plasma testosterone and luteinizing hormone levels in Figures 5C and 5D, and morphology of testis and prostate tissue for single Cdk8 knockout in Supplementary Figure 1C (although in this case the images do not appear very comparable between control and CDK8 KO, thus perhaps wider fields should be shown), but, for example, there is no analysis of different meiotic stages or of gene expression in single knockouts. It is worth mentioning that single knockouts seem to show a corresponding upregulation of the level of the paralogue kinase, indicating that any lack of phenotypes might be due to feedback compensation, which would be an interesting finding if confirmed; this has not been mentioned.

We agree that a description of the single KO could be beneficial, but we expect no big differences with the WT or Cre-Ert. We found neither histological differences nor changes in cell counts or ratios of cell types. Our ethical committee also has concerns about sacrificing mice without major phenotypic changes, without a well formulated hypothesis about the observed effects. We plan to add histological pictures to the next version of the article.

We have updated histological figures with new figures for iDKO and Cre+Tam mice with additional fields of view and better quality staining (2A-B).

The second major weakness is that the correlation between double knockout and reduced expression of genes involved in steroid hormone biosynthesis is portrayed as a causal mechanism for the phenotypes observed. While this is a possibility, there are no experiments performed to provide evidence that this is the case. Furthermore, there is no evidence showing that CDK8 and/or CDK19 are directly responsible for the transcription of the genes concerned.

We agree with the reviewer that the effects on CDK8/CDK19/CCNC could lead to the observed transcriptional changes in multiple indirect steps. There are, however, major technical challenges in examining the binding of transcription factors in the tissue, especially in Leydig cells which are a relatively minor population. We will clarify it in the revision and strengthen this point in the discussion.

We have added corresponding explanation in the Discussion: “We hypothesize that all these changes are caused by disruption of testosterone synthesis in Leydig cells, although, at this point, we cannot definitively prove that the affected genes are regulated by CDK8/19 directly.”

The claim of reproductive defects in the induced double knockout of CDK8/19 resulted from the loss of CCNC via a kinase-independent mechanism is interesting but was not supported by the data presented. While the construction and analysis of the systemic induced knockout model of Cdk8 in Cdk19KO mice is not trivial, the analysis and data are weakened by the systemic effect of Cdk8 loss, making it difficult to separate the systemic effect from the local testis effect.

We agree with the reviewer that the effects on the testes could be due to the systemic loss of CDK8 rather than specifically in the testis, and we will clarify it in the revision. We will also clarify that although our results are suggestive that the effects of CDK8/19 knockout are kinase-independent, and that the loss of Cyclin C is a likely explanation for the kinase independence, but we do not claim that it is *the* mechanism.

In this version we added several caveats indicating that the proposed mechanism is likely, but not the only one possible.

Also using TAM-treated wild type as control is ok, but a better control will be TAM-treated ERT2-cre; CDK8f/f or TAM-treated ERT2 Cre CDK19/19 KO, so as to minimize the impact from the well-recognized effect of TAM.

We used TAM-treated ERT2-cre for most of the experiments, and did not observe any major histological or physiological differences with the WT+TAM. We will make sure to present them in the revision.

The authors found that Sertoli cells re-entered the cell cycle in the inducible double knockout but stopped short of careful characterization other than increased expression of cell cycle genes.

Unfortunately, we were not able to perform satisfactory Ki67 staining to address this point.

Dko should be appropriately named iDKO (induced dKO). We will make the corresponding change.

We performed necropsy ? not the right wording here.

Colchicine-like apoptotic bodies ? what does this mean? Not clear.

We made appropriate changes - all DKO were renamed iDKO, necropsy changed to autopsy and cells designated as “apoptotic”.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Given the proprietary claims of the authors ("We have for the first time generated mice with the systemic inducible Cdk8 knockout on the background of Cdk19 constitutive knockout"), it does not appear acceptable and indeed might be misleading, to not describe the overall phenotypes of the mice. Are mice normal size/weight? Does an autopsy reveal anything other than atrophied genital tissue in males? Do the authors find a phenotype in the intestinal epithelium, as previously reported? (N.B. this could potentially clarify a discrepancy in the literature since the loss of the secretory lineages in double knockouts reported by the Firestein lab was not reproduced by intestinal organoid double knockout in the paper by the Fisher lab).

We have removed the statement “for the first time”, although to the best of our knowledge this is the fact. We did not attempt to describe all the phenotypic effects of the Cdk8/19 knockout in this paper, since some of the phenotypic observations related to mouse weight and behavior varied between different laboratories involved and require additional analysis. The effect on the urogenital system was by far the most striking histological feature observed and it was carefully addressed in this paper. Other findings require additional experiments and are out of the scope of this paper and we plan to focus on them later. As per suggestion of the reviewer we performed histological analysis of DKO intestines and found the same decrease in the Paneth and goblet cells numbers as described by Dannappel et al. We added corresponding figures (Supplemental fig. 1C) to the paper.

If the authors wish to reinforce their claims about causality of steroidogenic gene expression and phenotype, they could try rescuing the phenotype by treating mice with testosterone.

As stated in Discussion, we hypothesized that injection of testosterone would not rescue the phenotype, as the androgen receptor signaling is also affected. However we would like to perform such an experiment, but we were not able to procure testosterone pellets at this time.

If they wish to claim a direct effect of CDK8/19 on the expression of steroidogenic genes, they could also assess CDK8/19 binding to promoters of the genes analysed by ChIP.

There are big technical challenges in examining the binding of transcription factors in the primary tissue, especially in Leydig cells, a minor population, so we cannot perform such an experiment.

In order to conclude that their CDK8/19 inhibitor treatment worked, they could show target engagement by cell thermal shift assay, loss of CDK8/19 kinase-dependent gene expression, or loss of CDK8/19 substrate phosphorylation (eg interferon-induced STAT1 S727 phosphorylation) under the conditions used. Alternatively, they could show rescue with a kinase-dead allele.

As noted in public comments - we thank the reviewer for raising this concern. The target selectivity and target engagement by the inhibitors used in this study (Senexin B and SNX631-6) have been described in other models and published. CDK8/19 engagement and target selectivity of Senexin B, used in our vitro studies, have been extensively characterized in cell-based assays (Chen et al., Cells 2019, 8(11), 1413; Zhang et al., J Med Chem. 2022 Feb 24;65(4):3420-3433.) Similar characterization has been published for SNX631-6 and its equipotent analog SNX631, which showed drastic antitumor activity when used in vivo at the same dosing regimen as in this paper (Li et al., J Clin Invest. 2024;134(10):e176709). The comparison of the pharmacokinetics data obtained in the present study and in vitro activity of SNX631-6 in a cell-based assay suggests that the tissue concentrations of this drug should have provided substantial inhibition of Cdk8/19. Unfortunately, there are no known phosphorylation substrates specific for Cdk8/19 that can be used as pharmacodynamic markers. The widely used STAT1 phosphorylation at S727 is exerted not only by CDK8/19 but also by other kinases and shows variable response to CDK8/19 inhibition (Chen et al., Cells 2019, 8(11), 1413). In the revised MS, we have added a Western blot with pSTAT1 S727 staining of WT, 8KO, 19KO and iDKO testes. Cdk8/19 knockout did not decrease and apparently even increased the level of pSTAT1 S727, which demonstrates that this marker of CDK8/19 activity it is not suitable for our tissue type. While the evidence that Cdk8/19 kinase inhibition in the testes after in vivo drug treatment does not match the phenotype of iDKO is admittedly indirect, the same result has been obtained in the cell culture studies with Sertoli cells, where the inhibitor concentration (1 µM Senexin B) was much higher than needed for the maximal Cdk8/19 inhibition.

Finally, I did not find any legends to supplementary figures anywhere.

We apologize for not including legends for supplementary figures, and will correct that in the next version of the manuscript.

Additionally, we addressed the question about the sufficiency of the lipid supply for steroidogenesis in testes. There was a possibility that steroidogenesis is impossible due to the lack of cholesterol input, but OilRed staining revealed that the situation is the opposite: lipid content in iDKO testes is significantly higher than in WT testes. We added corresponding text to the article and the supplementary Fig. S6.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Paralogs CDK8 and CDK19 are regulatory kinases associated with the transcriptional Mediator complex. We have for the first time generated mice with the systemic inducible Cdk8 knockout on the background of Cdk19 constitutive knockout. Cdk8/19 double knockout (DKO) males, but not single Cdk8 and Cdk19 KO, had an atrophic reproductive system and were infertile. The DKO males lacked postmeiotic spermatids and spermatocytes after meiosis I pachytene. Testosterone levels were decreased whereas the amounts of the luteinizing hormone were unchanged. Single cell RNA sequencing showed marked differences in the expression of steroidogenic genes (such as Cyp17a1, Star and Fads) in Leydig cells concomitant with alterations in Sertoli cells and spermatocytes likely associated with impaired synthesis of steroids. Star and Fads were also downregulated in cultivated Leydig cells after DKO. The treatment of primary Leydig cells culture with a CDK8/19 inhibitor did not induce the same changes in gene expression as DKO, and prolonged treatment of mice with a CDK8/19 inhibitor did not affect the size of testes. DKO, in contrast to single knockouts or treatment with a CDK8/19 kinase inhibitor, led to depletion of cyclin C (CcnC), the binding partner of CDK8/19 that has been implicated in CDK8/19-independent functions. This suggests that the observed phenotype was likely mediated through kinase-independent activities of CDK8/19, such as CcnC stabilization.

Revised version: Paralogs CDK8 and CDK19 are regulatory kinases associated with the transcriptional Mediator complex. We have e generated mice with the systemic inducible Cdk8 knockout on the background of Cdk19 constitutive knockout. Cdk8/19 double knockout (iDKO) males, but not single Cdk8 and Cdk19 KO, had an atrophic reproductive system and were infertile. The iDKO males lacked postmeiotic spermatids and spermatocytes after meiosis I pachytene. Testosterone levels were decreased whereas the amounts of the luteinizing hormone were unchanged. Single cell RNA sequencing showed marked differences in the expression of steroidogenic genes (such as Cyp17a1, Star and Fads) in Leydig cells concomitant with alterations in Sertoli cells and spermatocytes likely associated with impaired synthesis of steroids. Star and Fads were also downregulated in cultivated Leydig cells after iDKO. The treatment of primary Leydig cells culture with a CDK8/19 inhibitor did not induce the same changes in gene expression as iDKO, and prolonged treatment of mice with a CDK8/19 inhibitor did not affect the size of testes. iDKO, in contrast to single knockouts or treatment with a CDK8/19 kinase inhibitor, led to depletion of cyclin C (CcnC), the binding partner of CDK8/19 that has been implicated in CDK8/19-independent functions. This suggests that the observed phenotype was likely mediated through kinase-independent activities of CDK8/19, such as CcnC stabilization.

Introduction

Earlier version: Multiple protein complexes regulate transcription, in response to developmental cues, changes in environment, damage and other factors, allowing for precise control of gene expression. Among such complexes the multiprotein Mediator complex is especially important for controlling transcription, regulated by a multitude of signaling pathways through their transcription factors (TFs). The Mediator complex bridges the regulatory enhancer regions with promoters thereby recruiting RNA polymerase II (RNA Pol II) and the preinitiation complex. The Mediator possesses its own regulatory subpart – the cyclin dependent kinase module (CKM), containing either the CDK8 or the CDK19 protein kinases, in complex with their binding partner cyclin C (CcnC), as well as proteins MED12 and MED13 (Luyties and Taatjes, 2022). While CKM is a nuclear complex that regulates transcription, both CcnC (Ježek et al., 2019) and MED12 (Zhang et al., 2020) were shown to have CKM-independent cytoplasmic activities. CDK8/19 kinase activity positively regulates transcription induced by different signals (Chen et al., 2023, 2017), at least in part by mediating the RNA Pol II release from pausing (Steinparzer et al., 2019). Besides, CDK8 and CDK19 directly phosphorylate a number of transcription factors such as STATs (Martinez-Fabregas et al., 2020), SMADs (Alarcón et al., 2009) and others, thereby modulating transcription of cognate genes. CDK8/19 kinase activity also negatively regulates transcription, at least, in part through downregulating the protein levels of all the components of the Mediator and CKM (Chen et al., 2023).

Revised version: Multiple protein complexes regulate transcription, in response to developmental cues, changes in environment, damage and other factors, allowing for precise control of gene expression. Among such complexes the multiprotein Mediator complex is especially important for controlling transcription, regulated by a multitude of signaling pathways through their transcription factors (TFs). The Mediator complex bridges the regulatory enhancer regions with promoters thereby recruiting RNA polymerase II (RNA Pol II) and the preinitiation complex. The Mediator possesses its own regulatory subpart – the cyclin dependent kinase module (CKM), containing either the CDK8 or the CDK19 protein kinases, in complex with their binding partner cyclin C (CcnC), as well as proteins MED12 and MED13 (Luyties and Taatjes 2022). While CKM is a nuclear complex that regulates transcription, both CcnC (Ježek et al. 2019) and MED12 (S. Zhang, O’Regan, and Xu 2020) were shown to have CKM-independent cytoplasmic activities. CDK8/19 kinase activity positively regulates transcription induced by different signals (Chen et al. 2017, 2023), at least in part by mediating the RNA Pol II release from pausing (Steinparzer et al. 2019). Besides, CDK8 and CDK19 directly phosphorylate a number of transcription factors such as STATs (Martinez-Fabregas et al. 2020), SMADs (Alarcón et al. 2009) and others, thereby modulating transcription of cognate genes. CDK8/19 kinase activity also negatively regulates transcription, at least, in part through downregulating the protein levels of all the components of the Mediator and CKM (Chen et al. 2023).

Earlier version: Despite important CDK8/19 functions discovered in primary cell culture and cell lines (Luyties and Taatjes, 2022), the insights gained from in vivo models are limited. Mice with constitutively knocked out Cdk19 (CDK19KO) were generated as a part of the COMP project and are basically asymptomatic, fertile, and have normal lifespan. It has been shown recently that Cdk19-/- Lin- hematopoietic stem cells divide slower (Zhang et al., 2022), however the blood cell count was unaffected. Heterozygous mice with constitutive Cdk8 knockout were asymptomatic but, when crossed, no homozygous pups were born (Postlmayr et al., 2020; Westerling et al., 2007). Moreover, conditional Cdk8 KO is almost asymptomatic in adult mice. Minor differences were observed in colon epithelial differentiation (Dannappel et al., 2022; Prieto et al., 2023) and tumorigenesis (McCleland et al., 2015) as well as in osteoclastogenesis (Yamada et al., 2022) after tissue specific Cdk8 knockout. Although two CDK8/19 inhibitors were reported to have systemic toxicity (Clarke et al., 2016), this toxicity was subsequently found to be due to off-target effects of these inhibitors (Chen et al., 2019), and several CDK8/19 inhibitors have reached clinical trials (clinicaltrials.gov NCT03065010, NCT04021368, NCT05052255, NCT05300438). Several studies reported the existence of kinase-independent phenotypic activities for both CDK8 (Kapoor et al., 2010; Menzl et al., 2019b) and CDK19 (Audetat et al., 2017; Steinparzer et al., 2019), but the only known biochemical activity of both Cdk8 and Cdk19 that is kinase-independent is the stabilization of their binding partner CcnC (Barette et al., 2001; Chen et al., 2023). Targeted degradation or knockout of both CDK8 and CDK19 dramatically reduces the cellular levels of CcnC, whereas the knockout of either Cdk8 or Cdk19 alone has little effect on CcnC levels (Chen et al., 2023). The effects of double knockout of both Cdk8 and Cdk19 on CcnC should therefore be considered in interpreting the effect of the double knockout.

Revised version: Despite important CDK8/19 functions discovered in primary cell culture and cell lines (Luyties and Taatjes 2022), the insights gained from in vivo models are limited. Mice with constitutively knocked out Cdk19 (CDK19 KO) were generated as a part of the COMP project and are basically asymptomatic, fertile, and have normal lifespan. It has been shown recently that Cdk19−/− Lin- hematopoietic stem cells divide slower (Z. Zhang et al. 2022), however the blood cell count was unaffected. Heterozygous mice with constitutive Cdk8 inducible knockout (CDK8 iKO) were asymptomatic but, when crossed, no homozygous pups were born (Postlmayr, Dumeau, and Wutz 2020; Westerling, Kuuluvainen, and Mäkelä 2007). Moreover, conditional CDK8 iKO is almost asymptomatic in adult mice. Minor differences were observed in colon epithelial differentiation (Dannappel et al. 2022; Prieto et al. 2023) and tumorigenesis (McCleland et al. 2015) as well as in osteoclastogenesis (Yamada et al. 2022) after tissue specific Cdk8 knockout. Although two CDK8/19 inhibitors were reported to have systemic toxicity (Clarke et al. 2016), this toxicity was subsequently found to be due to off-target effects of these inhibitors (Chen et al. 2019), and several CDK8/19 inhibitors have reached clinical trials (clinicaltrials.gov NCT03065010, NCT04021368, NCT05052255, NCT05300438). Several studies reported the existence of kinase-independent phenotypic activities for both CDK8 (Menzl et al. 2019; Kapoor et al. 2010) and CDK19 (Audetat et al. 2017; Steinparzer et al. 2019), but the only known biochemical activity of both CDK8 and CDK19 that is kinase-independent is the stabilization of their binding partner CcnC (Barette et al. 2001; Chen et al. 2023). Targeted degradation or knockout of both CDK8 and CDK19 dramatically reduces the cellular levels of CcnC, whereas the knockout of either Cdk8 or Cdk19 alone has little effect on CcnC levels (Chen et al. 2023). The effects of double knockout of both Cdk8 and Cdk19 on CcnC should therefore be considered in interpreting the effect of the double knockout.

Earlier version: We have, for the first time, generated mice with a conditional knockout of the Cdk8 gene on the constitutive Cdk19 KO background (CDK8/19 double knockout, DKO). We have found that DKO males were completely infertile, had an undersized and dedifferentiated reproductive system. In DKO males, spermatogenesis was blocked after pachytene of meiosis I. DKO showed down-regulation of key steroid pathway genes, such as Cyp17a1, Star and Lcn2, Sc5d, Fads2 and reduced production of testosterone. Expression of key meiotic genes was also deregulated in meiosis I spermatocytes, and the progression of meiosis was halted. These results indicated a key role of the CDK8/19-CcnC complex in the maintenance of male reproductive system.

Revised version: We have now generated mice with a conditional knockout of the Cdk8 gene on the constitutive Cdk19 KO background (CDK8/19 inducible double knockout, iDKO). We have found that iDKO males were completely infertile and had an undersized and dedifferentiated reproductive system. In iDKO males, spermatogenesis was blocked after pachytene of meiosis I. iDKO showed down-regulation of key steroid pathway genes, such as Cyp17a1, Star and Lcn2, Sc5d, Fads2 and reduced production of testosterone. Expression of key meiotic genes was also deregulated in meiosis I spermatocytes, and the progression of meiosis was halted. These results indicated a key role of the CDK8/19/CcnC complex in the maintenance of male reproductive system.

Results

Removed in the revised version: Cdk8/19 DKO mice are viable and lack Cyclin C

Removed in the revised version: Previously we have crossed Cdk8fl/fl mice with Rosa-Cre/ERT2 with tamoxifen-inducible Cre activity and demonstrated effective KO in all tissues except for ovaries and uterus (Ilchuk et al., 2022) (Fig. 1A). However, the CDK8 inactivation in the male reproductive system (testes) was efficient at the protein and DNA levels (Fig. 1B, Supplementary Fig. S1A). To investigate the effects of the double knockout of Cdk8 and Cdk19, we crossed Cdk8fl/fl/Rosa-Cre/ERT2 mice with Cdk19-/- mice (Zhang et al., 2022) [MGI:5607862, Cdk19em1(IMPC)J], which are asymptomatic and are maintained as homozygotes. We injected tamoxifen into 2 months old Cdk8fl/fl/Cdk19-/-/Rosa-Cre/ERT2 male mice. The DKO mice appeared to be viable. Age matched tamoxifen-treated wild-type mice served as a control.

Removed in the revised version: Cdk8/19 knockout blocks spermatogenesis in mice. (A) Crossing of Cdk8fl/fl, Cdk19-/- and Cre/ERT2 mice and formation of experimental (Cdk8fl/flCdk19-/-Cre/ERT2 + tamoxifen, Cdk8fl/flCre/ERT2 + tamoxifen and Cdk19-/-) and control (Cdk8fl/flCdk19-/-Cre/ERT2without tamoxifen and wild-type + tamoxifen) groups. (B) Confirmation of tamoxifen-induced CDK8KO in testes by Western blot. (C) Cyclin C protein is absent in DKO, but not single KO in the testes. (D) Time course of experiments. CDK8KO was activated by tamoxifen administration in males of 8-10 weeks old. Urogenital abnormalities became visible in two weeks. Spermatogenesis was analyzed by flow cytometry and immunofluorescence (IF) after 2, 8 and 20 weeks since activation. Single cell RNA sequencing was performed after 7 weeks of KO. (E) Male urogenital system atrophy in DKO mice. (F) H&E staining of prostate, epididymis and testes of DKO mice and tamoxifen-treated control. 100X magnification (G) H&E staining of WT and DKO seminiferous tubules. Yellow arrows indicate colchicine-like figures in the nuclei of dying cells.1000X magnification. (H) Sexual behavior and fertility of tamoxifen treated control, single KOs and DKO male mice.

Removed in the revised version: We have also analyzed the levels of the CDK8/19 binding partner CcnC after single and double KOs in mouse embryonic fibroblasts and in mouse testes after DKO. In agreement with previous observations in human cell lines (Chen et al., 2023), CcnC was undetectable in DKO embryonic fibroblasts as well as in DKO testes, but not in the single CDK19KO (Supplementary Fig. S1B, Fig. 1C). Hence, the stabilization of CcnC by CDK8 and CDK19 occurs not only in human but also in mouse cells and tissues.

Earlier version: Spermatogenesis is blocked in DKO males

Revised version: Spermatogenesis is blocked in CDK8/19 iDKO males

Added in the revised version: Previously we have crossed Cdk8fl/fl mice with Rosa-Cre/ERT2 with tamoxifen-inducible Cre activity (Fig. 1A, Supplementary Fig S1A) and demonstrated effective KO in all tissues except for ovaries and uterus (Ilchuk et al. 2022). However, the CDK8 inactivation in the male reproductive system (testes) was efficient at both genomic and protein levels (Supplementary Fig. S1B, Fig. 1B). To investigate the effects of the double knockout of Cdk8 and Cdk19, we crossed Cdk8fl/fl/Rosa-Cre/ERT2 mice with Cdk19−/− mice (Z. Zhang et al. 2022) [MGI:5607862, Cdk19em1(IMPC)J]. All the substrains were then maintained as homozygotes. We injected tamoxifen into 2 months old Cdk8fl/fl/Cdk19−/−/Rosa-Cre/ERT2 male mice. 2 months after tamoxifen injection iDKO mice had significantly lower body weight than tamoxifen injected control mice (Fig. 1C). The most noticeable difference was reduction of the male reproductive system - testes, epididymis and prostate (Fig 1D). Changes observed in iDKO intestines resembled those detected in Cdk8iIEC-KO/Cdk19−/− (Dannappel et al. 2022): number of Paneth cells and goblet cells was significantly decreased (Supplementary Fig.S1C). Age matched tamoxifen-treated Rosa-Cre/ERT2 mice served as a control.

Added in the revised version: Changes of male urogenital system in Cdk8/19 knockout mice.

Added in the revised version: (A) Crossing of Cdk8fl/fl, Cdk19−/− and Cre/ERT2 mice and formation of experimental (Cdk8fl/fl/Cdk19−/−/Rosa-Cre/ERT2 + tamoxifen, Cdk8fl/fl/Rosa-Cre/ERT2 + tamoxifen and Cdk19−/−) and control (Cdk8fl/fl/Cdk19−/−/Rosa-Cre/ERT2 without tamoxifen and wild-type + tamoxifen) groups. (B) Confirmation of tamoxifen-induced CDK8 iKO in testes by Western blot. (C) After 2 months of KO induction iDKO mice had significantly lower body weight. (D) Male urogenital system atrophy in iDKO mice.

Removed in the revised version: To investigate DKO effects we performed necropsy with subsequent H&E staining 8 weeks post tamoxifen treatment (Fig. 1D). The most striking results were observed in the urogenital system. The prostate, testes and the epididymis were significantly smaller compared to single KO and tamoxifen-treated wild-type mice (Fig. 1E), the diameter of testicular tubules was diminished (Fig. 1F). The germinal epithelium of the tubules was presented by Sertoli cells with typical large light nuclei and some spermatogonia or spermatocytes. Cells at the postmeiotic stages of spermatogenesis were completely absent. Among the spermatocytes there were only early prophase cells. In contrast, we detected cells with colchicine-like apoptotic bodies (Fig. 1F,G, yellow arrows). Thus, in these mice the spermatogenesis progenitor cells were present but their differentiation and/or meiosis was blocked. In Sertoli cells the vacuoles are the sign of the loss of contact between these cells and spermatogenic elements (Fig. 1F,G). Leydig cells were significantly smaller and were almost depleted of secretory vacuoles suggesting the lack of hormonal activity. The epididymis contained empty tubules; the epithelium appeared undifferentiated, that is, neither typical borders between epithelial cells nor clear Golgi zone were visible. The prostate and the epididymis were also atrophic (Fig. 1F). These changes were detectable from 2 weeks after tamoxifen treatment and persistent for at least 6 months (Fig. 1D). However, they were not observed in single KOs (Supplementary Fig. S1C).

Removed in the revised version: DKO and CDK8KO mice are infertile

Earlier version: To prove the DKO infertility we assessed their sexual behavior and number of pups they fathered. As tamoxifen has an impact on the male reproductive system (Willems et al., 2011), the experiments were performed at least 6 weeks post injection and we used tamoxifen-treated C57BL/6 mice as a control. Four groups were enrolled in the experiment: tamoxifen treated C57BL/6 (WT+Tam), tamoxifen treated Cdk8fl/flRosa-Cre/ERT2 (CDK8KO), Cdk8fl/flCdk19-/- Rosa-Cre/ERT2 (CDK19KO) and tamoxifen treated Cdk8fl/flCdk19-/-Rosa-Cre/ERT2 (DKO). Three males in each group were kept separately with two outbred CD1 females each for 3 months. Copulative plugs and the number of pups were checked 5 days a week (Fig. 1H). Tamoxifen-treated wild-type mice demonstrated normal fertility: 10 plugs were detected and 86 pups were born. The CDK19KO group showed slightly reduced fertility: 32 pups along with an increased number of the plugs (16) probably caused by the lower rate of pregnancy onset. Surprisingly, only two plugs and no pups were observed in the CDK8KO group despite normal appearance and behavior of males, and no microscopic abnormalities in the urogenital system (Supplementary Fig S1). Neither plugs nor pups were detectable in the DKO cohort (Fig. 1H). These experiments showed that both CDK8KO and DKO are infertile, and the cause of CDK8KO infertility is likely to be the lack of sexual behavior.

Revised version: Next, we decided to assess if the morphological changes in the reproductive system in the iDKO result in changes in their sexual behavior or number of pups they fathered. As tamoxifen has an impact on the male reproductive system (Willems et al. 2011), the experiments were performed at least 6 weeks post injection and we used tamoxifen-treated C57BL/6 mice as a control. Four groups were enrolled in the experiment: tamoxifen treated C57BL/6 (WT+Tam), tamoxifen treated Cdk8fl/flRosa-Cre/ERT2 (CDK8 iKO), Cdk8fl/flCdk19−/−Rosa-Cre/ERT2 (CDK19 KO) and tamoxifen treated Cdk8fl/flCdk19−/−Rosa-Cre/ERT2 (iDKO). Three males in each group were kept separately with two outbred CD1 females each for 3 months. Copulative plugs and the number of pups were checked 5 days a week (Fig. 1E). Tamoxifen-treated wild-type mice demonstrated normal fertility: 10 plugs were detected and 86 pups were born. The CDK19 KO group showed slightly reduced fertility: 32 pups along with an increased number of the plugs (16) probably caused by the lower rate of pregnancy onset. Surprisingly, only two plugs and no pups were observed in the CDK8 iKO group despite normal appearance and behavior of males. Neither plugs nor pups were detectable in the iDKO cohort (Fig. 1E). These experiments showed that both CDK8 iKO and iDKO are infertile, and the cause of CDK8 iKO infertility is likely to be the lack of sexual activity.

Added in the revised version: We injected tamoxifen into 2 months old Cdk8fl/fl/Cdk19−/−/Rosa-Cre/ERT2 male mice. Age matched tamoxifen-treated wild-type mice served as a control. To investigate iDKO effects we performed an autopsy with subsequent H&E staining 8 weeks post tamoxifen treatment (Fig. 2A). One of the morphological observations in iDKO mice was in the intestines, resembling changes observed in iDKO intestines resembled those detected in Cdk8iIEC-KO/Cdk19−/− (Dannappel et al. 2022): number of Paneth cells and goblet cells was significantly decreased (Supplementary Fig.S1C). However, the most noticeable difference between iDKO and control mice was reduction of the male reproductive system - testes, epididymis and prostate (Fig 1D). The prostate, testes and the epididymis were significantly smaller compared to single KOs and tamoxifen-treated wild-type mice (Fig. 2A), the diameter of testicular tubules was reduced (Fig. 2B).

Added in the revised version: Cdk8/19 knockout blocks spermatogenesis in mice.

Added in the revised version: (A) H&E staining of prostate, epididymis and testes of iDKO mice and tamoxifen-treated control. 100X magnification (B) H&E staining of WT and iDKO seminiferous tubules. 400X magnification. (G) Sexual behavior and fertility of tamoxifen treated control, single KOs and iDKO male mice. (C) Time course of experiments. CDK8 iKO was activated by tamoxifen administration in males of 8-10 weeks old. Urogenital abnormalities became visible in two weeks. Spermatogenesis was analyzed by flow cytometry and immunofluorescence (IF) after 2, 8 and 20 weeks since activation. Single cell RNA sequencing was performed after 7 weeks of KO. (D) Western blot analysis of Cre/ERT2 (Cre+Tam), single (CDK8 iKO and CDK19 KO) and double (iDKO) knockout testes 2 months after tamoxifen injections. CcnC protein is absent in iDKO, but not single KO in the testes. pSTAT727 is independent of CDK8/19 KO. Stars mark nonspecific staining by CDK19 antibodies. (E) CDK8 and CDK19 IHC staining of testes sections, 200x (upper row) and 630x magnification (bottom row) showing staining in various types of testicular cells. (F-K). Flow cytometry analysis of CDK8/CcnC expression in different testicular cell types. Figures F and I show major CDK8 (50.68%) and CcnC populations (44.52%), figures G and J show that 1n (round, but not elongated spermatids), 2n and 4n cells can be CDK8 and CcnC positive, figures H and K indicate, not only cKit+ cells among 2n-4n can be CDK8 and CcnC positive.

Added in the revised version: The germinal epithelium of the tubules was presented by Sertoli cells with typical large light nuclei and some spermatogonia or spermatocytes. Cells at the postmeiotic stages of spermatogenesis were completely absent. Among the spermatocytes there were only early prophase cells. In contrast, we detected an increased number of apoptotic cells (Fig. 2B). Thus, in these mice the spermatogenesis progenitor cells were present but their differentiation and/or meiosis was blocked. In Sertoli cells the vacuoles are the sign of the loss of contact between these cells and spermatogenic elements (Fig. 2A,B). Leydig cells were significantly smaller and were almost depleted of secretory vacuoles suggesting the lack of hormonal activity. The epididymis contained empty tubules; the epithelium appeared undifferentiated, that is, neither typical borders between epithelial cells nor clear Golgi zone were visible. The prostate and the epididymis were also atrophic (Fig. 2A). These changes were detectable from 2 weeks after tamoxifen treatment and persistent for at least 6 months (Fig. 2C). However, suprisingly they were not observed in CDK8 iKO (Supplemental Fig. S1D), which despite no microscopic changes, fathered no offspring and had no sexual behaviour (Fig. 1E).

Added in the revised version: CDK8/19 distribution in testes

Added in the revised version: Spermatogenesis failure in iDKO can be caused both by events in testes and extratesticular regulation. Internal events require expression of CDK8/19 at least in some testes cells and efficient KO of CDK8 which was not achieved in the female reproductive system (Ilchuk et al. 2022). Both CDK8 and CDK19 were present in testes and knockout was efficient after 7 days of daily treatment with 2 mg of tamoxifen (Fig. 2D). We have also analyzed the levels of the CDK8/19 binding partner CcnC after single and double KOs in mouse embryonic fibroblasts (MEFs) (Supplemental Fig. S1E) and in mouse testes after iDKO (Fig. 2D). In agreement with previous observations in human cell lines (Chen et al. 2023), CcnC was present if either CDK8 or CDK19 were present, but undetectable in iDKO embryonic fibroblasts as well as in iDKO testes. As CDK8 and CDK19 functions are considered to significantly overlap (Chen et al. 2023) it is possible that in the absence of one of the proteins the level of the other will increase blurring possible single KO phenotype. Indeed, we observed such compensation for MEFs, but not for testes (Fig. 2D and Supplemental Fig. S1E) indicating tissue specificity of this mechanism. Decrease in phosphorylation of STAT1 serine 727 is often used as CDK8/19 inhibition marker (Rzymski et al. 2017) despite the fact that different kinases can phosphorylate it (Chen et al. 2019). However, pSTAT727 level was not decreased in single or double KOs (Fig. 2D) indicating that pSTAT727 is not a suitable marker for this model.

Added in the revised version: Different testicular cell types contribute to the successful spermatogenesis and little is known about CDK8/19 expression in each type. McCleland et.al. performed IHC on different tissues and found CDK8 only in spermatogonia of all testicular cell types (McCleland et al. 2015). Our IHC analysis showed that CDK8 and CDK19 are present in several types of testicular cells, including Leydig cells, and cells inside the tubule. Interestingly CDK8 and CDK19 appear to be expressed in different types of cells in the seminiferous tubules, with CDK8 mostly expressed in the periphery of the tubule and in lower levels in the center of the tubule, while CDK19 was mostly expressed in the center of the tubule (Fig. 2E). We also used flow cytometry to identify cell types expressing CDK8/19. To detect cells with the expression of at least one kinase we used anti-CcnC antibodies. In wild type male mice CDK8 (Fig. 2F-H) and CcnC (Fig. 2I-K) positive cells were detected among 1n (round, but not elongated spermatids), 2n and 4n cells, indicating their role in all spermatogenesis stages. In agreement with McCleland’s report all c-Kit+ cells were positive for CcnC and CDK8, but other 2n/4n c-Kit negative cells were also positive for both proteins (Fig. 2H, K).

Earlier version: Spermatogenic cells of DKO mice are unable to advance through meiosis I prophase

Revised version: Spermatogenic cells of iDKO mice are unable to advance through meiosis I prophase

Removed in the revised version: To obtain a quantitative spermatogenesis pattern we performed cell cycle analysis by flow cytometry with propidium iodide staining. We examined wild-type, single KOs and DKO animals sacrificed two months after tamoxifen treatment. Cell cycle analysis revealed striking differences in DKO mice, with disappearance of elongated spermatids, almost all round spermatids and massive cell death (Fig. 2A,B). All other groups had the same normal cell distribution (Fig. 2B).

Removed in the revised version: Absence of postmeiotic 1n cells in DKO 2 months after KO induction. (A) Distinctive histograms of wild-type (left) and DKO (right) mice. Violet - 4n population, red - 2n population, green - round spermatids, blue - elongated spermatids, orange - apoptotic subG1 cells. (B) Quantitative distribution of testes cells between these groups. Wild type with and without tamoxifen as control groups have similar distribution to CDK8 and CDK19 single KO. DKO testes have greatly reduced number of round spermatids and no elongated spermatids. (C) Overall cellularity is also significantly reduced only in DKO testes. (D) IF staining of control and DKO testes frozen sections. Nuclei are stained by DAPI (blue pseudocolor), SYCP3 is depicted as green, γH2A.X - as red. All stages of spermatogenesis are visible in control testes, while pachytene is the last detected stage in DKO. Confocal microscopy, magnification 600X.

Earlier version: Noticeably, the 4n population and cells in the S-phase were not impaired by DKO indicating that DKO spermatogonia successfully entered meiosis but could not produce haploid spermatids. This ongoing “meiotic catastrophe” led to almost full depopulation of the testes in DKO mice (Fig. 2С). Due to fast transition from meiosis I to meiosis II it is difficult to detect secondary spermatocytes by flow cytometry. To specify the stage of CDK8/19-dependent spermatogenesis failure we performed phospho-γH2A.X and SYCP3 immunofluorescent staining of seminiferous tubules. Phospho-γH2A.X is a DNA damage marker which marks double-strand breaks in leptotene, zygotene and in sex chromosomes during the pachytene. SYCP3 is a component of the synaptonemal complex, expressed during meiosis I prophase, which forms distinct patterns in the late zygotene/pachytene. In the wild-type tubules all the meiosis stages were visible (Fig. 2D). At the same time almost all meiotic cells in DKO were blocked in the pachytene (Fig. 2D), indicating that DKO cells enter meiosis but cannot traverse through meiosis I prophase and subsequently undergo cell death.

Revised version: To obtain a quantitative spermatogenesis pattern we performed cell cycle analysis by flow cytometry with propidium iodide staining. We examined wild-type, single KOs and iDKO animals sacrificed two months after tamoxifen treatment. Cell cycle analysis revealed striking differences in iDKO mice, with disappearance of elongated spermatids, almost all round spermatids and massive cell death (Fig. 3A, B). All other groups had the same normal cell distribution (Fig. 3B). Noticeably, the 4n population and cells in the S-phase were not impaired by iDKO indicating that iDKO spermatogonia successfully entered meiosis but could not produce haploid spermatids. This ongoing “meiotic catastrophe” led to almost full depopulation of the testes in iDKO mice (Fig. 3C). Due to fast transition from meiosis I to meiosis II it is difficult to detect secondary spermatocytes by flow cytometry. To specify the stage of CDK8/19-dependent spermatogenesis failure we performed phospho-γH2A.X and SYCP3 immunofluorescent staining of seminiferous tubules. Phospho-γH2A.X is a DNA damage marker which marks double-strand breaks in leptotene, zygotene and in sex chromosomes during the pachytene. SYCP3 is a component of the synaptonemal complex, expressed during meiosis I prophase, which forms distinct patterns in the late zygotene/pachytene. In the wild-type tubules all the meiosis stages were visible (Fig. 3D). At the same time almost all meiotic cells in iDKO were blocked in the pachytene (Fig. 3D), indicating that iDKO cells enter meiosis but cannot traverse through meiosis I prophase and subsequently undergo cell death.

Added in the revised version: Absence of postmeiotic 1n cells in DKO 2 months after KO induction.

Added in the revised version: (A) Distinctive histograms of wild-type (left) and DKO (right) mice. Violet - 4n population, red - 2n population, green - round spermatids, blue - elongated spermatids, orange - apoptotic subG1 cells. (B) Quantitative distribution of testes cells between these groups. Wild type with and without tamoxifen as control groups have similar distribution to CDK8 and CDK19 single KO. iDKO testes have greatly reduced number of round spermatids and no elongated spermatids. (C) Overall cellularity is also significantly reduced only in DKO testes. (D) IF staining of control and DKO testes frozen sections. Nuclei are stained by DAPI (blue pseudocolor), SYCP3 is depicted as green, γH2A.X - as red. All stages of spermatogenesis are visible in control testes, while pachytene is the last detected stage in DKO. Confocal microscopy, magnification 600X. M - meiotic entry spermatocytes; L - leptotene; Ph - pachytene; PM - post-meiotic stages.

Single cell RNA sequencing

Earlier version: To investigate molecular mechanisms of CDK8/19 mediated alterations in meiosis we performed single cell RNA sequencing (scRNAseq) of the testes. We analyzed cell suspensions from testes of two tamoxifen-treated wild-type С57/BL/6 and two DKO animals.

Revised version: To investigate molecular mechanisms of CDK8/19 mediated alterations in meiosis we performed single cell RNA sequencing (scRNAseq) of the testes. We analyzed cell suspensions from testes of two tamoxifen-treated wild-type C57/BL/6 and two iDKO animals.

Earlier version: The analysis of cell type composition confirmed the reduced number of post pachytene spermatocytes and almost complete absence of spermatids while meiotic entry, leptotene and zygotene cell numbers remained unaffected (Fig. 3A,B). At the same time the percentage of undifferentiated spermatogonial stem cells in DKO animals was unchanged, nor was the proportion of Leydig cells altered. Sertoli cells became the most abundant cell type in DKO mice with reduced testes cellularity, however their absolute number didn’t change significantly.

Revised version: The analysis of cell type composition confirmed the reduced number of post pachytene spermatocytes and almost complete absence of spermatids while meiotic entry, leptotene and zygotene cell numbers remained unaffected (Fig. 4A,B). At the same time the percentage of undifferentiated spermatogonial stem cells in iDKO animals was unchanged, nor was the proportion of Leydig cells altered. Sertoli cells became the most abundant cell type in iDKO mice with reduced testes cellularity, however their absolute number didn’t change significantly.

Added in the revised version: Single cell RNA sequencing reveals loss of spermatids due to steroidogenesis failure.

Earlier version: Single cell RNA sequencing reveals loss of spermatids due to steroidogenesis failure. (A) UMAP projection and relative cell numbers for all testicular cell types in control and DKO samples. Number of secondary spermatocytes is significantly decreased and spermatids are almost absent in DKO samples. (B) UMAP projection and relative cell numbers for spermatogonia and primary spermatocytes. Post-pachytene spermatocytes are severely depleted in DKO samples. (C) GO Biological Processes pathways enriched among Leydig cells differentially expressed genes (DEGs). Lipid metabolism and steroid biosynthesis is severely perturbed. (D) Violin plots for key Leydig cells’ genes.

Revised version: (A) UMAP projection and relative cell numbers for all testicular cell types in control and iDKO samples. Number of secondary spermatocytes is significantly decreased and spermatids are almost absent in iDKO samples. (B) UMAP projection and relative cell numbers for spermatogonia and primary spermatocytes. Post-pachytene spermatocytes are severely depleted in iDKO samples. (C) GO Biological Processes pathways enriched among Leydig cells differentially expressed genes (DEGs). Lipid metabolism and steroid biosynthesis is severely perturbed. (D) Violin plots for key Leydig cells genes.

Differential gene expression

Earlier version: Differentiation of sperm lineage is a tightly controlled process, where changes in functioning in a certain type of cells can affect the viability and differentiation of other cells. To understand the underlying mechanism of meiotic cell death in DKOs we compared gene expression in each type of cell clusters present both in wild-type and knockout animals.

Revised version: Differentiation of sperm lineage is a tightly controlled process, where changes in functioning in a certain type of cells can affect the viability and differentiation of other cells. To understand the underlying mechanism of meiotic cell death in iDKOs we compared gene expression in each type of cell clusters present both in wild-type and knockout animals.

Leydig cells

Earlier version: Leydig cells are the primary source for testosterone in the testes, a hormone required for spermatogenesis. scRNAseq analysis showed significant transcriptomic changes in Leydig cells. Among 126 genes that were differentially expressed (|log2FC|>0.4; p<0.01) in WT vs DKO Leydig cells, 30 were down-regulated, however among 14 strongly affected genes with |log2FC|>1 11 were down-regulated. Strikingly, 9 of 14 strongly downregulated genes were associated with lipid (specifically steroid) metabolism (Supplementary table S1) and 3 of these 9 genes were linked to male infertility (Aherrahrou et al., 2020; Song, 2007; Stoffel et al., 2008).

Revised version: Leydig cells are the primary source for testosterone in the testes, a hormone required for spermatogenesis. scRNAseq analysis showed significant transcriptomic changes in Leydig cells. Among 126 genes that were differentially expressed (|log2FC|>0.4; p<0.01) in WT vs iDKO Leydig cells, 30 were down-regulated, however among 14 strongly affected genes with |log2FC|>1 11 were down-regulated. Strikingly, 9 of 14 strongly downregulated genes were associated with lipid (specifically steroid) metabolism (Supplemental table S1) and 3 of these 9 genes were linked to male infertility (Aherrahrou et al. 2020; Stoffel et al. 2008; W.-C. Song 2007).

Earlier version: According to the GO analysis (Fig. 3C, Supplementary table S2), the most prominent changes were downregulation of lipid metabolism and the steroid hormone biosynthesis. The most significantly down-regulated gene was Cyp17a1, showing 10-fold decrease (Fig. 3D). The Cyp17a1 gene in male mice is mainly expressed in Leydig cells (Missaghian et al., 2009). The product of this gene catalyzes the key reaction of the steroidogenic pathway. Mice with constitutive Cyp17a1 KO are phenotypically female (Aherrahrou et al., 2020). Downregulated Lcn2 (Fig. 3D) specifically expressed in Leydig cells codes for a protein of the lipocalin family that transports small hydrophobic molecules including steroids and is affected by fertility manipulations (Kang et al., 2017; Yanai et al., 2021). Another downregulated gene, Fads2 (Fig. 3D), is a key enzyme for biosynthesis of highly unsaturated fatty acids. Its KO causes infertility in males by arresting the spermatogonial cycle at the stage of the round spermatids (Stroud et al., 2009). The Sult1e1 (Fig. 3D) gene encodes sulfotransferase which inactivates estrogens and its deficiency leads to hyperplasia of Leydig cells and atrophy of seminiferous tubules (Song, 2007). The Star, Sc5d, Pld3 and Apoc1 (Fig. 3D) genes, also involved in lipid metabolism, were down-regulated more than 2-fold. Interestingly, the Hsd3b6 (Fig. 3D) is a surprisingly upregulated gene associated with steroid metabolism in Leydig cells.

Revised version: According to the GO analysis (Fig. 4C, Supplemental table S2), the most prominent changes were downregulation of lipid metabolism and the steroid hormone biosynthesis. The most significantly down-regulated gene was Cyp17a1, showing 10-fold decrease in expression (Fig. 4D). The Cyp17a1 gene in male mice is mainly expressed in Leydig cells (Missaghian et al. 2009). The product of this gene catalyzes the key reaction of the steroidogenic pathway. Mice with constitutive Cyp17a1 KO are phenotypically female (Aherrahrou et al. 2020). Downregulated Lcn2 (Fig. 4D) specifically expressed in Leydig cells codes for a protein of the lipocalin family that transports small hydrophobic molecules including steroids and is affected by fertility manipulations (Kang et al. 2017; Yanai et al. 2021). Another downregulated gene, Fads2 (Fig. 4D), is a key enzyme for biosynthesis of highly unsaturated fatty acids. Its KO causes infertility in males by arresting the spermatogonial cycle at the stage of the round spermatids (Stroud et al. 2009). The Sult1e1 (Fig. 4D) gene encodes sulfotransferase which inactivates estrogens and its deficiency leads to hyperplasia of Leydig cells and atrophy of seminiferous tubules (W.-C. Song 2007). The Star, Sc5d, Pld3 and Apoc1 (Fig. 3D) genes, also involved in lipid metabolism, were down-regulated more than 2-fold. Interestingly, the Hsd3b6 (Fig. 4D) is a surprisingly upregulated gene associated with steroid metabolism in Leydig cells.

Earlier version: Not only testosterone production was decreased in double knockouts. Expression of Insl3 gene, coding for an important peptide hormone secreted by Leydig cells (Esteban-Lopez and Agoulnik, 2020), was also reduced (Supplementary table S3).

Revised version: Not only testosterone production was decreased in double knockouts. Expression of Insl3 gene, coding for an important peptide hormone secreted by Leydig cells (Esteban-Lopez and Agoulnik 2020), was also reduced (Supplemental table S3).

Earlier version: Among up-regulated genes, were several transcription factors governing steroidogenesis in Leydig cells (de Mattos et al., 2022; Ye et al., 2017). One of them, Nr2f2 (Fig. 3D) encodes COUP-TFII which regulates genes immediately involved in steroidogenesis (Cyp17a1, Hsd3b1, Cyp11a1 and Akr1c14) alongside several other Leydig-specific genes (Insl3, Amhr2) (Fig. 3D). Another gene, Cebpb, encodes TF C/EBPβ which activates Star transcription and stimulates expression of Nr4a1 (also known as Nur77) which is itself a key regulator of steroidogenesis in Leydig cells.

Revised version: Among up-regulated genes, there were several transcription factors governing steroidogenesis in Leydig cells (Ye et al. 2017; de Mattos, Viger, and Tremblay 2022). One of them, Nr2f2 (Fig. 4D) encodes COUP-TFII which regulates genes immediately involved in steroidogenesis (Cyp17a1, Hsd3b1, Cyp11a1 and Akr1c14) alongside several other Leydig-specific genes (Insl3, Amhr2) (Fig. 3D). Another gene, Cebpb, encodes TF C/EBPβ which activates Star transcription and stimulates expression of Nr4a1 (also known as Nur77) which is itself a key regulator of steroidogenesis in Leydig cells.

The Kit gene (down-regulated), encoding a tyrosine kinase receptor, and Kitl gene (up-regulated), encoding its ligand, were also present among differentially expressed genes. Both of these genes play an important role in spermatogenesis (knockouts and mutants are infertile), Leydig cell maturation and steroidogenesis (Liu et al. 2017; Ye et al. 2017).

We hypothesize that the observed upregulation of several steroidogenesis regulators may be a part of a compensation feedback loop in response to the decreased testosterone level.

Sertoli cells

Earlier version: ScRNAseq (Fig. 3A, B) shows that Sertoli cell fraction is greatly increased in DKO specimens, however this is mostly due to the reduction of germ cell number. At the same time, analysis of cell cycle-specific genes revealed that in DKO Sertoli cells can re-enter the cell cycle (Fig. 4A). In particular, genes involved in G0-S transition and G2-M transition were upregulated.

Revised version: ScRNAseq (Fig. 4A) shows that Sertoli cell fraction is greatly increased in iDKO specimens, however this is mostly due to the reduction of germ cell number. At the same time, analysis of cell cycle-specific genes revealed that in iDKO Sertoli cells can re-enter the cell cycle. In particular, genes involved in G0-S transition and G2-M transition were upregulated (Fig. 5A).

Added in the revised version: DKO Sertoli cells re-enter the cell cycle and lose characteristic cytoskeleton organization.

Earlier version: DKO Sertoli cells re-enter the cell cycle and lose characteristic cytoskeleton organization. (A) Violin plots for Reactome cell cycle gene sets indicate that Sertoli cells in DKO lose terminal differentiation and re-enter cell cycle. Percentage of cells in G1-S and G2-M transitions are increased in DKOs. (B) Violin plots for key cytoskeleton and intercellular contacts related DEGs. (C) IF staining for vimentin demonstrates blood-testis barrier (BTB) integrity disruption and loss of characteristic striation cytoskeleton patterns in DKOs. Magnification 600X. (D) Enrichment of GO stress pathways in Sertoli cells indicates their dysfunction in DKOs.

Revised version: (A) Violin plots for Reactome cell cycle gene sets indicate that Sertoli cells in iDKO lose terminal differentiation and re-enter cell cycle. Percentage of cells in G1-S and G2-M transitions are increased in iDKOs. (B) Violin plots for key cytoskeleton and intercellular contacts related DEGs. (C) IF staining for vimentin demonstrates blood-testis barrier (BTB) integrity disruption and loss of characteristic striation cytoskeleton patterns in DKOs. Magnification 600X. (D) Enrichment of GO stress pathways in Sertoli cells indicates their dysfunction in iDKOs.

Earlier version: According to downregulation of the steroid hormone biosynthesis in Leydig cells we expected to see in Sertoli cells patterns similar to those in mice with Sertoli cell specific KOs of androgen receptor (AR) – SCARKO (Sertoli Cells Androgen Receptor KnockOut) (De Gendt et al., 2014; Larose et al., 2020) and luteinizing hormone (LH) receptor – LURKO (Griffin et al., 2010). However, it must be noted that these mice have constitutive receptor KOs and do not fully develop their reproductive system, whereas our KO model mice develop normally and become fertile before the knockout induction.

Revised version: According to downregulation of the steroid hormone biosynthesis in Leydig cells we expected to see in Sertoli cells patterns similar to those in mice with Sertoli cell specific KOs of androgen receptor (AR) – SCARKO (Sertoli Cells Androgen Receptor KnockOut) (Larose et al. 2020; De Gendt et al. 2014) and luteinizing hormone (LH) receptor – LURKO (Griffin et al. 2010). However, it must be noted that these mice have constitutive receptor KOs and do not fully develop their reproductive system, whereas our KO model mice develop normally and become fertile before the knockout induction.

Earlier version: Among DEGs we identified a number of upregulated and downregulated genes consistent with their changes in SCARKO (Defb45, Myl6, Tmsb4x, Espn, Ldhb among top DEGs). At the same time, we compared our DEG set with those published earlier (De Gendt et al., 2014; Larose et al., 2020) and observed little overlap between our data and published results as well as between the two published papers themselves (Supplementary Fig. S3).

Revised version: Among DEGs we identified a number of upregulated and downregulated genes consistent with their changes in SCARKO (Defb45, Myl6, Tmsb4x, Espn, Ldhb among top DEGs). At the same time, we compared our DEG set with those published earlier (Larose et al. 2020; De Gendt et al. 2014) and observed little overlap between our data and published results as well as between the two published papers themselves (Supplemental Fig. S3).

However, a much better agreement was found between our data and the study by deGendt et.al. (De Gendt et al. 2014) at the level of GO Molecular Function and Cellular Compartment categories.

Earlier version: According to the GO Cellular Compartment (Supplementary table S2), both up- and down-regulated genes are associated with cytoskeleton (Vim, Actb, Actg1) (Show et al., 2003), apical-basal axis (Cdc42) (Heinrich et al., 2021), cellular membrane (Atp1a1, Gja1) (Rajamanickam et al., 2017; Sridharan et al., 2007) and intercellular contacts (Espn, Anxa2, S100a10) (Chojnacka et al., 2017; Willems et al., 2010). These closely related structures play an important role in Sertoli cell functioning: formation of the blood-testicular barrier (BTB) and intercommunication with developing germ cells (Fig. 4B) (Wong and Cheng, 2009). Therefore, our findings on differential gene expression in Sertoli cells are in agreement with data obtained in mice with impaired AR signaling. Among the top enriched biological processes several up-regulated stress response pathways were also found, indicating deterioration of the Sertoli cells in the absence of testosterone (Fig. 4D).

Revised version: According to the GO Cellular Compartment (Supplemental table S2), both up- and down-regulated genes are associated with cytoskeleton (Vim, Actb, Actg1) (Show et al. 2003), apical-basal axis (Cdc42) (Heinrich et al. 2021), cellular membrane (Atp1a1, Gja1) (Rajamanickam, Kastelic, and Thundathil 2017; Sridharan et al. 2007) and intercellular contacts (Espn, Anxa2, S100a10) (Willems et al. 2010; Chojnacka, Bilinska, and Mruk 2017). These closely related structures play an important role in Sertoli cell functioning: formation of the blood-testicular barrier (BTB) and intercommunication with developing germ cells (Fig. 5B) (Wong and Cheng 2009). In agreement with this, immunostaining also revealed perturbed vimentin polymerization (Fig. 5C). Therefore, our findings on differential gene expression in Sertoli cells are in agreement with data obtained in mice with impaired AR signaling. Among the top enriched biological processes several up-regulated stress response pathways were also found, indicating deterioration of the Sertoli cells in the absence of testosterone (Fig. 5D).

This molecular evidence of Sertoli cell spatial organization disturbance and disruption of cell contacts is in good agreement with the patterns observed during histological examination.

Germ cells

Earlier version: As primary spermatocytes of double knockout mice stop their differentiation in pachytene, we compared gene expression in two groups of germ cells: undifferentiated spermatogonia and meiosis-entry (leptotene-zygotene-pachytene) spermatocytes. 1107 DEGs were identified for undifferentiated cells, 1650 for spermatocytes, and 874 of these genes were common. Almost all the DEGs were up-regulated (99.2%). This overlap may be explained by the fact that undifferentiated and differentiated spermatogonia represent a continuous spectrum with gradually changing transcription patterns. Unique genes for the Early Spermatogonia cluster were attributed by GO as genes related to RNA processing and biosynthesis and ribosome biogenesis, whereas no stem cell specific pathways were found. We conclude that Cdk8/19 knockout and disruption of steroid biosynthesis and testosterone production have no significant impact specifically on early spermatogonia.

Revised version: As primary spermatocytes of double knockout mice stop their differentiation in pachytene, we compared gene expression in two groups of germ cells: undifferentiated spermatogonia and meiosis-entry (leptotene-zygotene-pachytene) spermatocytes. 1107 DEGs were identified for undifferentiated cells, 1650 for spermatocytes, and 874 of these genes were common. Almost all the DEGs were up-regulated (99.2%) (Supplementary Fig. S4). This overlap may be explained by the fact that undifferentiated and differentiated spermatogonia represent a continuous spectrum with gradually changing transcription patterns. Unique genes for the Early Spermatogonia cluster were attributed by GO as genes related to RNA processing and biosynthesis and ribosome biogenesis, whereas no stem cell specific pathways were found. We conclude that CDK8/19 knockout and disruption of steroid biosynthesis and testosterone production have no significant impact specifically on early spermatogonia.

Earlier version: Most cells in knockout testes do not progress to stages after pachytene. Accordingly, apoptosis-related pathways, stress response pathways and the autophagy pathway were upregulated (Supplementary table S2, Supplementary Fig. S4). More specifically, a set of meiosis-specific genes was also up-regulated. These genes were related to chromosome pairing, DNA recombination and histone modification, which are known to be involved in meiosis. We hypothesize that this effect might be caused by compensatory mechanisms triggered by inability to proceed through meiosis.

Revised version: Most cells in knockout testes do not progress to stages after pachytene. Accordingly, apoptosis-related pathways, stress response pathways and the autophagy pathway were upregulated (Supplemental table S2, Supplementary Fig. S5). More specifically, a set of meiosis-specific genes was also up-regulated. These genes were related to chromosome pairing, DNA recombination and histone modification, which are known to be involved in meiosis. We hypothesize that this effect might be caused by compensatory mechanisms triggered by inability to proceed through meiosis.

Validation of single cell RNA data

To further prove the steroid hypothesis we performed CYP17A1 immunostaining of testicular cross sections and western blotting, as well as direct testosterone measurement in the blood via LC-MS and cell culture experiments with ex vivo Leydig cells.

Earlier version: In control testes, CYP17A1 protein was localized outside of seminiferous tubules, in Leydig cells. In DKO animals, CYP17A1 staining was completely absent (Fig. 5A), despite Leydig cells being present in histological sections (Fig. 1G). WB (Fig. 5B) further confirming this finding.

Revised version: In control testes, CYP17A1 protein was localized outside of seminiferous tubules, in Leydig cells. In iDKO animals, CYP17A1 staining was completely absent (Fig. 6A), despite Leydig cells being present in histological sections (Fig. 2A,B). Western blot analysis of all genotypes further confirmed this finding (Fig 6B).

Added in the revised version: Confirmation of scRNA sequencing data by other methods.

Earlier version: Confirmation of single cell RNA sequencing data with other methods. (A) IF staining for CYP17A1 of testes frozen sections, magnification 600X. CYP17A1 is visualized in extratubular space in Leydig cells in control mice and is completely absent in DKOs. (B) Western blot for CYP17A1 confirms disappearance of the protein in DKOs, but not in other genotypes. (C) In agreement with this testosterone blood level is decreased only in DKO mice. (D) Luteinizing hormone production is not impaired by CDK8KO and CDK19KO or DKO.

Revised version: (A) IF staining for CYP17A1 of testes frozen sections, magnification 600X. CYP17A1 is visualized in extratubular space in Leydig cells in control mice and is completely absent in iDKOs. (B) Western blot for CYP17A1 confirms disappearance of the protein in iDKOs, but not in other genotypes. (C) Testosterone blood level is decreased only in iDKO mice. (D) Luteinizing hormone production is not impaired by CDK8 iKO and CDK19 KO or iDKO.

Earlier version: CYP17A1 catalyzes several key reactions leading to transformation of pregnenolone into testosterone. We measured the concentration of the key androgen hormone - testosterone in the blood of wild-type mice, single KOs and DKO. Only DKO showed significantly lower levels of testosterone, consistent with only this genotype presenting abnormal spermatogenesis (Fig. 5C).

Revised version: CYP17A1 catalyzes several key reactions leading to transformation of pregnenolone into testosterone. We measured the concentration of the key androgen hormone - testosterone in the blood of wild-type mice, single KOs and iDKO. Only iDKO showed significantly lower levels of testosterone, consistent with only this genotype presenting abnormal spermatogenesis (Fig. 6C).

Earlier version: Changes in steroid synthesis in Leydig cells can be caused by several mechanisms – specifically CDK8/19/Cyclin C can regulate CYP17A1 expression in Leydig cells themselves, or CDK8/19/Cyclin C may be involved in regulation of brain hormonal signals, which control testosterone production. To test the latter hypothesis we measured the concentration of luteinizing hormone in blood serum. The results showed no difference in LH level through all the genotypes (Fig. 5D). This means that CDK8/19/Cyclin C acts directly in Leydig cells to regulate steroidogenesis.

Revised version: Changes in steroid synthesis in Leydig cells can be caused by several mechanisms. Specifically CDK8/19/CcnC can regulate CYP17A1 expression in Leydig cells themselves, or CDK8/19/CcnC can be involved in regulation of brain hormonal signals, which control testosterone production. To test the latter hypothesis we measured the concentration of luteinizing hormone in blood serum. The results showed no difference in LH level through all the genotypes (Fig. 6D). Therefore it is more plausible that CDK8/19/CcnC acts directly in Leydig cells to regulate steroidogenesis.

Added in the revised version: CDK8 was also shown to positively regulate lipid metabolism in D. melanogaster (Tang et al. 2018; Xiao Li et al. 2022). In this case lack of steroidogenesis could be explained by lack of its precursor cholesterol. We detected OilRed positive stained cells in the intertubular space in good agreement with Leydig cells localisation. Moreover, iDKO staining was more intense than in wild type mice, indicating accumulation of lipids, which are not converted to testosterone (Supplementary Fig. S6).

Earlier version: Spermatogenesis in DKO mice slightly recovers with time

Revised version: Spermatogenesis in iDKO mice slightly recovers with time

Earlier version: To address the question of the persistence of spermatogenesis failure we performed histological and flow cytometry analysis of mice 5 months after tamoxifen-treatment. A certain percentage of 1n cells (primarily round spermatids) have reappeared 3-5 months after tamoxifen treatment (Fig. 6A), nevertheless the total number of cells and 1n cells remained low compared to control (Fig. 6B). H&E staining of prostate and epididymis also revealed slight alleviation of the phenotype, but no mature sperm (Fig. 6С). Immunofluorescent staining of seminiferous tubules confirmed this finding demonstrating solitary tubules with postmeiotic cells (Fig. 6D). However, CYP17A1 was not restored in Leydig cells (Fig. 6E).

Revised version: To address the question of the persistence of spermatogenesis failure we performed histological and flow cytometry analysis of mice 5 months after tamoxifen-treatment. A certain percentage of 1n cells (primarily round spermatids) have reappeared 3-5 months after tamoxifen treatment (Fig. 7A), nevertheless the total number of cells and 1n cells remained low compared to control (Fig. 7B). H&E staining of prostate and epididymis also revealed slight alleviation of the phenotype, but no mature sperm (Fig. 7C). Immunofluorescent staining of seminiferous tubules confirmed this finding demonstrating solitary tubules with postmeiotic cells (Fig. 7D). However, the levels of CYP17A1 were not restored in Leydig cells (Fig. 7E).

Added in the revised version: Limited recovery of spermatogenesis 5 months after iDKO induction.

Earlier version: А minor recovery of spermatogenesis 5 months after DKO induction. (A) Round and elongated spermatids are once again detected by flow cytometry 5 months after DKO, however, (B) overall testes cellularity is only slightly increased. (C) Postmeiotic cells become visible at H&E staining of the tubules, however, epididymal ducts remain empty. (D) Post-pachytene and post-meiotic (PM) cells became visible on the SYCP3 + γH2A.X stained frozen sections, magnification 600X. (E) CYP17A1 level remains at the background level 5 months after KO induction.

Revised version: (A) Round and elongated spermatids become detectable by flow cytometry 5 months after iDKO. (B) Overall testes cellularity is only slightly increased. (C) Postmeiotic cells become visible at H&E staining of the tubules, however, epididymal ducts remain empty. (D) Post-pachytene and post-meiotic (PM) cells became visible on the SYCP3 + γH2A.X stained frozen sections, magnification 600X. (E) CYP17A1 level remains at the background level 5 months after KO induction.

Pharmacological inhibition of CDK8/19 does not affect spermatogenesis

Earlier version: Recent findings have established that CDK8/19 may have kinase-independent functions (Steinparzer et al., 2019) including stabilization of Cyclin C by CDK8/19 (Chen et al., 2023). Depletion of CDK8/19 in testes leads to CCNC degradation (Fig. 1C). To examine if the observed phenotype is related to the kinase activity of CDK8/19 or is kinase-independent we treated WT mice with a pharmacological CDK8/19 inhibitor – SNX631-6 (Li et al., 2023). 11 weeks old male C57BL/6 mice were treated with SNX631-6 medicated chow (500 ppm, 40-60 mg/kg/day dosage on average) or control diet for 3 weeks. No abnormal clinical observations were identified in treated mice through the whole treatment period. Testes were weighed and fixed in 10% formalin for H&E staining. There were no significant differences in testes weight between control and treated group and histology analysis did not reveal any noticeable difference (Fig. 7A,B), in contrast to the DKO testes. However, the drug concentration in the testes measured by LCMS/MS was lower than in blood, suggesting the effects of the blood-testis barrier, but still ∼10 times higher than the drug’s IC50 in the cell-based assay - 50 ng/ml. (Fig. 7C).

Revised version: Recent findings have established that CDK8/19 may have kinase-independent functions (Steinparzer et al. 2019) including stabilization of CcnC in a kinase independent manner (Chen et al. 2023). Depletion of CDK8/19 in testes of iDKO mice leads to CcnC degradation (Fig. 2D). To examine if the observed phenotype is related to the kinase activity of CDK8/19 or is kinase-independent we treated WT mice with a pharmacological CDK8/19 inhibitor – SNX631-6 (J. Li et al. 2024). 11 weeks old male C57BL/6 mice were treated with SNX631-6 medicated chow (500 ppm, 40-60 mg/kg/day dosage on average) or control diet for 3 weeks. The same SNX631-6 dosing regimen was highly efficacious in suppressing castration resistant prostate cancer growth in male mice (J. Li et al. 2024). No abnormal clinical observations were identified in treated mice through the whole treatment period. Testes were weighed and fixed in 10% formalin for H&E staining. There were no significant differences between control and treated group in histology analysis or in testes weight (Fig. 8A,B), in contrast to the iDKO testes. The drug concentration in the testes measured by LCMS/MS was lower than in blood (Fig. 8C), suggesting the effects of the blood-testis barrier, but this concentration (50 ng/ml) was still ∼10 times higher than the drug’s IC50 in the cell-based assay - 50 ng/ml.

Earlier version: Effects of CDK8/19 inhibitor on spermatogenesis in mice (A-C) male C57BL/6 mice were treated with SNX631-6 medicated chow (500 ppm, 40-60 mg/kg/day dosage on average) for three weeks. (A) Representative images of H&E histology analysis of the testes tissues collected from animals in control or treated groups. (B) organ weights of testes (left and right) at endpoint. (C) SNX631-6 concentrations in plasma and testes tissues at endpoint. (D) qPCR analysis of steroidogenic Star and Fads genes in ex vivo cultured Leydig cells in response to CDK8/19i Senexin B (1 μM) or hydroxytamoxifen-induced CDK8/19 DKO.

Revised version: Effects of CDK8/19 inhibitor on spermatogenesis in mice (A-C) male C57BL/6 mice were treated with SNX631-6 medicated chow (500 ppm, 40-60 mg/kg/day dosage on average) for three weeks.

Added in the revised version: (A) Representative images of H&E histology analysis of the testes tissues collected from animals in control or treated groups. (B) Organ weights of testes (left and right) at endpoint. (C) SNX631-6 concentrations in plasma and testes tissues at endpoint. (D) qPCR analysis of steroidogenic Star and Fads genes in ex vivo cultured Leydig cells in response to CDK8/19 inhibitor Senexin B (1 μM) or hydroxytamoxifen-induced CDK8/19 iDKO.

Earlier version: The effects of DKO and CDK8/19 kinase inhibition were compared in ex vivo primary Leydig cell culture from WT mice and Cdk8fl/flCdk19-/-Cre-Rosa mice. We treated the first with CDK8/19 inhibitor Senexin B and the latter with hydroxytamoxifen to induce Cdk8-/- in vitro. Then we compared expression of Cyp17a1, Star and Fads genes downregulated according to scRNA sequencing data in treated and untreated cells (Fig. 7D). Star and Fads were downregulated in OH-Tam treated cells, whereas basal Cyp17a1 expression level was too low to make any conclusion. At the same time, CDK8/19 inhibition didn’t affect expression levels of Star and Fads, suggesting that the effects of the DKO were kinase-independent.

Revised version: The effects of iDKO and CDK8/19 kinase inhibition were compared in ex vivo primary Leydig cell culture from WT mice and Cdk8fl/flCdk19−/−Cre-Rosa mice. We treated the first with CDK8/19 inhibitor Senexin B and the latter with hydroxytamoxifen to induce Cdk8−/− in vitro. Then we compared expression of Cyp17a1, Star and Fads genes that were downregulated according to scRNA sequencing data in treated and untreated cells (Fig. 8D). Star and Fads were downregulated in OH-Tam treated cells, whereas basal Cyp17a1 expression level was too low to make any conclusion. At the same time, CDK8/19 kinase inhibitor did not affect expression levels of Star and Fads, suggesting that the effects of the iDKO were kinase-independent.

Discussion

Earlier version: The Mediator complex is a key part of transcriptional regulation through signaling pathways. The enzymatic component of the Mediator-associated CKM module – CDK8 and its paralog CDK19, have emerged as co-regulators in several signal-activated transcriptional pathways, such as Wnt, TGF-beta, androgen and estrogen signaling, STATs, response to serum growth factors and others (reviewed in (Menzl et al., 2019a)). Despite that, the knockout of Cdk8 is only critical in embryogenesis (Westerling et al., 2007), with no phenotype in adult organisms (McCleland et al., 2015), while Cdk19 is wholly dispensable. The same tendency is observed for the knockout of MED subunits and other CKM proteins: constitutive KOs are often embryonically lethal (except for paralog of MED12 - MED12L), while tissue-specific knockouts in adult animals rarely have a severe phenotype (Ilchuk et al., 2023).

Revised version: The Mediator complex is a key part of transcriptional regulation through signaling pathways. The enzymatic components of the Mediator-associated CKM module – CDK8 and its paralog CDK19, have emerged as co-regulators in several signal-activated transcriptional pathways, such as Wnt, TGF-beta, androgen and estrogen signaling, STATs, response to serum growth factors and others (reviewed in (Menzl, Witalisz-Siepracka, and Sexl 2019)). Despite that, the knockout of Cdk8 is only critical in embryogenesis (Westerling, Kuuluvainen, and Mäkelä 2007), with no phenotype in adult organisms (McCleland et al. 2015), while CDK19 is wholly dispensable. The same tendency is observed for the knockout of MED subunits and other CKM proteins: constitutive KOs are often embryonically lethal (except for MED12L a paralog of MED12), while tissue-specific knockouts of these genes in adult animals rarely have a severe phenotype (Ilchuk et al. 2023).

Added in the revised version: Hypothetically CDK8 and CDK19 can compensate for each other in KO animals and mitigate the phenotype in single KOs. To confirm this, we produced the first ubiquitous knockout of Cdk8/19 in 2 months old animals. In addition to the specific phenotype described here, the analysis of tissues from the knockout animals provided the first in vivo confirmation of two key molecular observations on the general CDK8/19 activities that were previously described only in vitro and that are not yet widely understood as potential caveats in CDK8/19 studies. The first observation (Chen et al. 2023) is that the knockout of CDK8 and CDK19 leads to the loss of CcnC, in contrast to CDK8/19 kinase inhibition that actually increases the protein levels of CcnC. In the present study, we found that CcnC is depleted both in MEF and in the testes of iDKO mice. This finding underscores that CDK8/19 kinase inhibition may have very different phenotypic consequences from the knockout of these genes. The second finding was that STAT1 phosphorylation at serine 727, a widely used pharmacodynamic marker of CDK8/19 kinase activity, shows variable response to CDK8/19 inhibition, which depends on the cell type and the duration of CDK8/19 inhibitor treatment, and that it can be induced by different signals in a CDK8/19-independent manner (Chen et al. 2019). Here we found that serine 727 phosphorylated STAT1 was undiminished (and possibly even increased) by Cdk8/19 knockout in the testes of iDKO mice (Fig. 2D). Hence, a reliance on STAT1 phosphorylation as a pharmacodynamic marker of CDK8/19 activity can be misleading.

Added in the revised version: Previously described animals with conditional knockout of Cdk8 in the intestine on Cdk19−/− background presented with a mild phenotype, moderately affecting the number of specialized secreting cells (Dannappel et al. 2022). This phenotype was confirmed by histological observations in our model. The most prominent phenotype in our iDKO mice, evidenced at the physiological, morphological and molecular levels, was found in the testes, matching the finding that the testes express the highest levels of CDK8 and CDK19 RNA among all the normal human tissues (J. Li et al. 2024). In particular, we observed the lack of haploid spermatids and mature spermatozoa (Fig. 2A-B), with primary spermatocytes blocked in pachytene of meiosis I (Fig. 3). CDK8 and CDK19 were expressed in different types of testicular cells, putatively identified as spermatogonia (high expression), Leydig cells, and Seroli cells, and spermatocytes (weak expression) for CDK8, and Sertoli cells and spermatocytes (high expression) for CDK19. CDK8 and CDK19 were co-localized in some cells, although cells with high expression of CDK8 had weaker CDK19 staining and vice versa (Fig 2E). Flow cytometry analysis also showed high levels of CDK8 and CCNC in 2n, 4n and in some haploid round spermatids (but not elongated spermatids) (Fig. 2F-G, I-J). We also confirmed previous reports of McCleland et. al. identifying CDK8 expression in c-Kit+ spermatogonia cells (McCleland et al. 2015), but also show expression of CDK8 and CCNC in other 2n and 4n c-Kit-cells (Fig. 2H, K). Surprisingly, CDK8 iKO males previously described as asymptomatic revealed their infertility due to the lack of sexual behavior, despite the absence of morphological changes in the urogenital system (Fig. 1E, Supplemental Fig. S1D).

Removed in the revised version: Hypothetically CDK8 and CDK19 can compensate for each other in KO animals and mitigate the phenotype in single KOs. To confirm this we produced the first ubiquitous knockout of Cdk8/19 in 2 months old animals. Previously described animals with conditional knockout of Cdk8 in the intestine on Cdk19-/- background presented with a mild phenotype, moderately affecting the number of specialized secreting cells (Dannappel et al., 2022). Our DKO mice had a phenotype evident at the histological level: they lacked haploid spermatids and mature spermatozoa (Fig. 1F), with primary spermatocytes blocked in pachytene of meiosis I (Fig. 2). Surprisingly, Cdk8-/- males previously described as asymptomatic revealed their infertility due to the lack of sexual behavior, despite the absence of morphological changes in the urogenital system (Fig. 1E).

Earlier version: Mice with Med1fl/fl/Vasa-Cre tissue specific KO had an opposite phenotype to the one described here: prophase I passage was accelerated in these mice and spermatocytes entered zygotene/pachytene stage prematurely (Huszar et al., 2015). The difference between these effects is quite interesting in the context of the concept that association with CKM attenuates Mediator-dependent transcription either by sterically preventing TF binding (Freitas et al., 2022), or by post-transcriptional regulation of Mediator complex subunits (Chen et al., 2023; Poss et al., 2016). As far as we know, besides the current work, Huszar et al., 2015 is the only paper establishing a connection between any components of the Mediator complex and spermatogenesis.

Revised version: Mice with Med1fl/fl/Vasa-Cre tissue specific KO had an opposite phenotype to the one described here: prophase I passage was accelerated in these mice and spermatocytes entered zygotene/pachytene stage prematurely (Huszar et al. 2015). The difference between these effects is quite interesting in the context of the concept that association with CKM attenuates Mediator-dependent transcription either by sterically preventing TF binding (Freitas et al. 2022), or by post-transcriptional regulation of Mediator complex subunits (Poss et al. 2016; Chen et al. 2023). As far as we know, besides the current work, Huszar et al., 2015 is the only paper establishing a connection between any components of the Mediator complex and spermatogenesis.

Earlier version: We demonstrated a disruption of steroid biosynthesis in Leydig cells which could play the key role for the observed phenotype. A number of works (reviewed in (Ilchuk et al., 2023)) affirm a significant role of Mediator complex in lipid metabolism and biosynthesis as one of the main cofactors for main adipogenic TFs such as C/EBPα, PPARγ and SREBP.

Revised version: We demonstrated a disruption of steroid biosynthesis in Leydig cells which could play the key role for the observed phenotype. A number of studies (reviewed in (Ilchuk et al. 2023)) affirm a significant role of Mediator complex in lipid metabolism and biosynthesis as one of the main cofactors for the key adipogenic TFs such as C/EBPα, PPARγ and SREBP.

Earlier version: We detected a decrease in steroidogenic gene expression by scRNAseq (primarily Cyp17a1), and then confirmed this finding by WB, immunofluorescence, and direct testosterone measurement. This corresponds well with perturbed Leydig cells morphology (reduced size and secretory vacuoles depletion) and explains observed atrophy of prostate and epididymis (Fig. 1F). It is worth mentioning that constitutive Cyp17a1 knockout in mice leads to female phenotype in XY mice (Aherrahrou et al., 2020). Besides Cyp17a1, several other steroid metabolism related genes were downregulated (Lcn, Fads, Star, Sc5d, Sult1e1). Interestingly, the main feature of Fads (Stoffel et al., 2008) and Sult1e1 (Song, 2007) KOs is male infertility, caused by spermatogenesis arrest.

Revised version: We detected a decrease in steroidogenic gene expression by scRNAseq (primarily Cyp17a1), and then confirmed this finding by WB, immunofluorescence, and direct testosterone measurement. This corresponds well with perturbed Leydig cells morphology (reduced size and secretory vacuoles depletion) and explains observed atrophy of prostate and epididymis (Fig. 2A). It is worth mentioning that constitutive Cyp17a1 knockout in mice leads to female phenotype in XY mice (Aherrahrou et al. 2020). Besides Cyp17a1, several other steroid metabolism related genes were downregulated (Lcn, Fads, Star, Sc5d, Sult1e1). Interestingly, the main feature of Fads (Stoffel et al. 2008) and Sult1e1 (W.-C. Song 2007) KOs is male infertility, caused by spermatogenesis arrest.

Earlier version: There is much evidence that meiosis in spermatocytes depends on AR signaling (Wang et al., 2022). Several papers established a connection between AR signaling, and the Mediator complex (Russo et al., 2019). MED1 and MED19 directly interact with the AR and regulate its transcriptional activity (Jin et al., 2012; Weber et al., 2021). Additional evidence of the role of the Mediator complex and CDK8/19 in the urogenital system comes from studies of their role in prostate cancer, where CDK19 is elevated together with the AR signaling (Becker et al., 2020).

Revised version: There is substantial evidence that meiosis in spermatocytes depends on AR signaling (Wang, Li, and Yang 2022). Several papers established a connection between AR signaling, and the Mediator complex (Russo, Nouri, and Balk 2019). MED1 and MED19 directly interact with the AR and regulate its transcriptional activity (Jin, Claessens, and Fondell 2012; Weber, Ruoff, and Garabedian 2021). Additional evidence of the role of the Mediator complex and CDK8/19 in the urogenital system comes from studies of their role in prostate cancer, where CDK19 is elevated together with the AR signaling (Becker et al. 2020) and where CDK8/19 inhibition and MED12 knockdown partially inhibit the transcriptional AR signaling (J. Li et al. 2024; Andolfi et al. 2024) and where CDK8/19 inhibition and MED12 knockdown partially inhibit the transcriptional AR signaling (J. Li et al. 2024; Andolfi et al. 2024).

Earlier version: The DKO has also dramatically impacted Sertoli cells. The SCARKO mice were engineered two decades ago and show morphological abnormalities (i.e, reduced testes, epididymis and prostate, meiotic arrest during meiosis I prophase) similar to DKO (Chang et al., 2004; De Gendt et al., 2004). The scRNAseq technique used to elucidate the molecular mechanism of the meiotic arrest revealed a number of similarities with DKO data for Sertoli cells, especially in terms of pathways (Cao et al., 2021). Thus, genes associated with cytoskeleton, cellular membrane and cell-cell contacts were enriched among the DEGs (Fig. 4B). Nevertheless, the overlap of genes affected in the SCARKO and DKO mice was low (Supplementary Fig. S3). However, an important difference between our model and SCARKO is that Cdk8 knockout is induced after puberty. Another important consideration is that Sertoli cells are not the only cell type impacting spermatogenesis through AR signaling. AR signaling in peritubular myoid cells is also essential for successful spermatogenesis (Welsh et al., 2009). In this regard, probably a more relevant model for comparison with our results is described in (Stanton et al., 2012) where androgen signaling was suppressed in grown-up rats by low dose testosterone, or its combination with an AR antagonist. In this paper the group with the most complete inhibition of AR signaling displayed increased apoptosis in spermatocytes and full blockage of second meiotic division.

Revised version: The iDKO has also dramatically impacted Sertoli cells. The SCARKO mice were engineered two decades ago and show morphological abnormalities (i.e, reduced testes, epididymis and prostate, meiotic arrest during meiosis I prophase) similar to iDKO (De Gendt et al. 2004; Chang et al. 2004). The scRNAseq technique used to elucidate the molecular mechanism of the meiotic arrest revealed a number of similarities with iDKO data for Sertoli cells, especially in terms of pathways (Cao et al. 2021). Thus, genes associated with cytoskeleton, cellular membrane and cell-cell contacts were enriched among the DEGs (Fig. 5B). Nevertheless, the overlap of genes affected in the SCARKO and iDKO mice was low (Supplementary Fig. S3). However, an important difference between our model and SCARKO is that Cdk8 knockout is induced after puberty. Another important consideration is that Sertoli cells are not the only cell type impacting spermatogenesis through AR signaling. AR signaling in peritubular myoid cells is also essential for successful spermatogenesis (Welsh et al. 2009). In this regard, probably a more relevant model for comparison with our results is described in (Stanton et al. 2012) where androgen signaling was suppressed in grown-up rats by low dose testosterone, or its combination with an AR antagonist. In that paper the rats with the most complete inhibition of AR signaling displayed increased apoptosis in spermatocytes and full blockage of second meiotic division.

Earlier version: Just as in (Stanton et al., 2012), genes differentially expressed in spermatocytes blocked in prophase I can be divided in the following groups: (1) genes with known roles in meiosis, especially DNA synapse, homologous recombination and DNA repair, (2) genes associated with cellular stress and apoptosis, and (3) genes with roles in RNA processing and splicing. Changes in the first group can be explained by the compensatory mechanisms and changes in the second group correlate well with the observed extensive cell death of spermatocytes, changes in the third group remain enigmatic.

Revised version: Just as in (Stanton et al. 2012), genes differentially expressed in spermatocytes blocked in prophase I can be divided in the following groups: (1) genes with known roles in meiosis, especially DNA synapse, homologous recombination and DNA repair, (2) genes associated with cellular stress and apoptosis, and (3) genes with roles in RNA processing and splicing. Changes in the first group can be explained by the compensatory mechanisms and changes in the second group correlate well with the observed extensive cell death of spermatocytes, whereas changes in the third group remain enigmatic.

Noteworthy, 99.2% of spermatocyte DEGs with |logFc|>0.4 were upregulated genes, which is very similar to the findings of Pelish et al. (Pelish et al. 2015) and agrees with the negative regulation of the Mediator complex by CDK8/19 (Chen et al. 2023).

Earlier version: Based on these results we can conclude that DKO causes perturbance in different testicular cell types: Leydig cells, Sertoli cells and spermatocytes. We hypothesize that all these changes are caused by disruption of testosterone synthesis in Leydig cells. Data for other cell types are consistent with this hypothesis, however there is still a possibility that CDK8/19 have distinct roles in other testicular cell types which are not visible against the background of the testosterone reduction. However, our data are not sufficient to make the conclusion about the role of CDK8/19/Cyclin C directly in meiosis, besides its hormonal regulation. To answer this question, definitive mouse strains with cell-type specific KO in Sertoli cells and meiotic spermatocytes would be needed.

Revised version: Based on these results we can conclude that iDKO causes perturbance in different testicular cell types: Leydig cells, Sertoli cells and spermatocytes. We hypothesize that all these changes are caused by disruption of testosterone synthesis in Leydig cells, although, at this point, we cannot definitively prove that the affected genes are regulated by CDK8/19 directly. Data for other cell types are consistent with this hypothesis, however there is still a possibility that CDK8/19/CcnC have distinct roles in other testicular cell types which are not visible against the background of the testosterone reduction. However, our data are not sufficient to make the conclusion about the role of CDK8/19/CncC directly in meiosis, besides its hormonal regulation. To answer this question, definitive mouse strains with cell-type specific KO in Sertoli cells and meiotic spermatocytes would be needed.

Earlier version: The effects of DKO on spermatogenesis were maximal at 1-2 months after tamoxifen administration. At later time points we detected partial restoration of spermatogenesis, and signs of re-differentiation of the epididymis and activity of secreting cells in the prostate. This restoration did not lead to the production of mature sperm, and the total number of haploid cells increased insignificantly. There are several possible explanations for such an effect. It is unlikely that restoration happened in cells, which did not activate Cre-Lox, as the absence of CDK8 has been confirmed by us through PCR (Supplementary Fig. S1A). Additionally, no restoration of CYP17A1 was found in 5-month-old animals. There is a possibility that transcription of genes previously regulated by CDK8/19/Cyclin C was rewired through other transcriptional regulators. An intriguing possibility is activation of a compensatory AR-independent pathway that was recently shown in a zebrafish model (Zhai et al., 2022). An argument against such hypothesis would be the absence of such compensation in Cyp17a1 KO animals or SCARKO mice, but in the former animals the male reproductive system does not develop at all, and in the latter, non-canonical AR signaling through Src kinase may lead to progression to round spermatid stage (Cooke and Walker, 2021).

Revised version: The effects of iDKO on spermatogenesis were maximal at 1-2 months after tamoxifen administration. At later time points we detected partial restoration of spermatogenesis, and signs of re-differentiation of the epididymis and activity of secreting cells in the prostate. This restoration did not lead to the production of mature sperm, and the total number of haploid cells increased insignificantly. There are several possible explanations for such an effect. It is unlikely that restoration happened in cells, which did not activate Cre-Lox, as the absence of CDK8 has been confirmed by us through PCR (Supplemental Fig. S1A). Additionally, no restoration of CYP17A1 was found in 5-month-old animals. There is a possibility that transcription of genes previously regulated by CDK8/19/CcnC was rewired through other transcriptional regulators. An intriguing possibility is activation of a compensatory AR-independent pathway that was recently shown in a zebrafish model (Zhai et al. 2022). An argument against such hypothesis would be the absence of such compensation in Cyp17a1 KO animals or SCARKO mice, but in the former animals the male reproductive system does not develop at all, and in the latter, non-canonical AR signaling through Src kinase may lead to progression to round spermatid stage (Cooke and Walker 2021).

Earlier version: The luteinizing hormone produced in the pituitary gland is a major regulator of steroidogenesis in Leydig cells. It was shown that LH regulates expression level of Cyp17a1 and several other steroidogenic genes (Li et al., 2021; Ma et al., 2004). As CDK8KO in our model is ubiquitous, infertility could be caused by decrease of LH level due to improper regulation in the brain. However, we measured LH level in the serum and found it unaffected by single or double KOs (Fig 5D). That means that CDK8/19/Cyclin C are required for steroidogenic genes’ transcription in Leydig cells. And indeed, in our ex vivo experiment with Leydig primary cell culture, even without addition of LH, expression of Star and Fads was downregulated upon in vitro induction of CDK8KO by hydroxytamoxifen on the CDK19KO background.

Revised version: The luteinizing hormone produced in the pituitary gland is a major regulator of steroidogenesis in Leydig cells. It was shown that LH regulates expression level of Cyp17a1 and several other steroidogenic genes (Xiaoheng Li et al. 2021; Ma et al. 2004). As CDK8 iKO in our model is ubiquitous, infertility could be caused by decrease of LH level due to improper regulation in the brain. However, we measured LH level in the serum and found it unaffected by single or double KOs (Fig 6D). That means that CDK8/19/CcnC are required for steroidogenic genes’ transcription in Leydig cells, although we cannot definitely prove that CDK8/19 directly regulates the affected genes. And indeed, in our ex vivo experiment with Leydig primary cell culture, even without addition of LH, expression of Star and Fads was downregulated upon in vitro induction of CDK8 iKO by hydroxytamoxifen on the CDK19 KO background.

Added in the revised version: Steroidogenesis decline could be related to the altered overall lipid metabolism and precursor deficiency. Several groups reported CDK8/19 role in adipogenesis (Xiao Li et al. 2022) and lipid biosynthesis (Tang et al. 2018). However, contrary to this, we detected high levels of lipids in iDKO Leydig cells (Supplemental Fig. S6). Similar effect was detected in StAR knockout mice (Hasegawa et al. 2000), further confirming that downregulation of steroidogenic gene transcription is the primary cause of steroidogenesis failure in CDK8/19 mice.

Earlier version: The results of the present study reveal for the first time the role of CDK8/19/Cyclin C in the function of the male reproductive system. The key question regarding the results described here is whether the phenotypic effects are due to the inhibition of CDK8/19 kinase activity or to the loss of Cyclin C, which is protected by CDK8 and CDK19 from proteasome-mediated degradation (Chen et al., 2023), or other kinase-independent functions of CDK8/19. Here we show that Cyclin C essentially disappears both from the embryo fibroblasts and from the testes of the DKO mice. Given the known CDK8/19-independent functions of Cyclin C (Ježek et al., 2019), and the effect of Cyclin C knockout on lipid accumulation (Song et al., 2022), the likelihood of Cyclin C rather than CDK8/19 mediating the observed effects should be investigated. As we have shown, 30-day treatment of 11 weeks old mice with a potent CDK8/19 kinase inhibitor produced no anatomical changes in the testes. Furthermore, ex vivo treatment of Leydig cells with a CDK8/19 inhibitor did not reproduce the DKO effect on gene expression, suggesting that the effects of DKO were kinase-independent and possibly mediated by Cyclin C. A definitive answer to the role of Cyclin C vs CDK8/19 kinase activity will require a phenotypic comparison of DKO mice with mice expressing kinase-inactive mutants of CDK8/19 or inducible Cyclin C knockout.

Revised version: The results of the present study reveal for the first time the role of CDK8/19/CcnC in the function of the male reproductive system. The key question regarding the results described here is whether the phenotypic effects are due to the inhibition of CDK8/19 kinase activity or to the loss of CcnC, which is protected by CDK8 and CDK19 from proteasome-mediated degradation (Chen et al. 2023), or other kinase-independent functions of CDK8/19. Here we show that CcnC essentially disappears both from the embryo fibroblasts and from the testes of the iDKO mice. Given the known CDK8/19-independent functions of CcnC (Ježek et al. 2019), and the effect of CcnC knockout on lipid accumulation (Z. Song et al. 2022), the likelihood of CcnC rather than CDK8/19 mediating the observed effects should be investigated. As we have shown, 30-day treatment of 11 weeks old mice with a potent CDK8/19 kinase inhibitor produced no anatomical changes in the testes, and previous reports examining CDK8/19 did not produce similar effects. Furthermore, ex vivo treatment of Leydig cells with a CDK8/19 inhibitor did not reproduce the iDKO effect on gene expression, suggesting that the effects of iDKO were kinase-independent and possibly mediated by CcnC. A definitive answer to the role of CcnC vs CDK8/19 kinase activity will require a phenotypic comparison of iDKO mice with mice expressing kinase-inactive mutants of CDK8/19 or inducible CcnC knockout.

Materials and Methods

Earlier version: Animals and conditional DKO

Revised version: Animals and conditional iDKO

Earlier version: Cdk8fl/fl and R26-Cre-ERT2 (Jax:008463; B6.129-Gt(ROSA)26Sortm1(cre/ERT2)Tyj/J) provenance and genotyping procedures are described (Ilchuk et al., 2022). Cdk19-/- (Cdk19em1(IMPC)J) were purchased from Jackson Laboratory. Mice were genotyped by real-time PCR using oligonucleotides listed in Supplementary table S4.

Revised version: Cdk8fl/fl and R26-Cre-ERT2 (Jax:008463; B6.129-Gt(ROSA)26Sortm1(cre/ERT2)Tyj/J) provenance and genotyping procedures are described (Ilchuk et al. 2022). Cdk19−/−, mutant mouse strain C57BL/6N-Cdk19em1(IMPC)J/Mmucd was obtained from the MMRRC at UC Davis KOMP, RRID:MMRRC_047035-UCD.

Earlier version: Animals were maintained under controlled room conditions (22–24°C and a 14 h light: 10 h dark photoperiod). Mice were given access to food and water ad libitum. All manipulations with animals were performed according to the Local Bioethical Committee recommendations and according to the Declaration of Helsinki (1996). Tamoxifen treatment was performed as described previously: 8-10 week old males were injected with 3 mg of tamoxifen daily for 7 consecutive days (Ilchuk et al., 2022). Tamoxifen used for KO induction is an estrogen receptor modulator and can affect the mouse organism, especially fertility and hormonal levels, per se (Willems et al., 2011). To ensure that observed effects are caused by the KOs and not by tamoxifen treatment we used tamoxifen-treated controls in every experiment and conducted all experiments but one (Fig. 1D) at least one month after tamoxifen treatment. Mice were sacrificed by cervical dislocation. To evaluate fertility, males were kept with two CD-1 outbred female mice each. Copulative plugs and newborn pups were monitored daily in the morning.

Revised version: Mice were genotyped by real-time PCR using oligonucleotides listed in Supplemental table S4. Animals were maintained under controlled room conditions (22–24°C and a 14 h light: 10 h dark photoperiod). Mice were given access to food and water ad libitum. All manipulations with animals were performed according to the Local Bioethical Committee recommendations and according to the Declaration of Helsinki (1996). Tamoxifen treatment was performed as described previously: 8-10 week old males were injected with 3 mg of tamoxifen daily for 7 consecutive days (Ilchuk et al. 2022). Tamoxifen used for KO induction is an estrogen receptor modulator and can by itself affect the mouse organism, especially fertility and hormonal levels (Willems et al. 2011). To ensure that the observed effects are caused by the KOs and not by tamoxifen treatment we used tamoxifen-treated controls in every experiment and conducted all experiments but one (Fig. 2C) at least one month after tamoxifen treatment. Mice were sacrificed by cervical dislocation. To evaluate fertility, males were kept with two CD-1 outbred female mice each. Copulative plugs and newborn pups were monitored daily in the morning.

Histology

Earlier version: For primary tissue examination organs were fixed in 10% formaldehyde, paraffin embedded, sectioned on a microtome and stained with H&E.

Revised version: For primary tissue examination organs were fixed in 10% formaldehyde, paraffin embedded, sectioned on a microtome and stained with H&E. We used Periodic acid-Shiff (PAS) staining to highlight molecules with a high percentage of carbohydrate content to identify Paneth cells and goblet cells in the ileum.

Added in the revised version: For immunofluorescence analysis paraffin sections, the sections were deparaffinized and hydrated using the following steps: 10 min in xylene (twice), 7 min in 100% ethanol (twice), 7 min in 95% ethanol (twice), and 5 min in water at room temperature (three times). Antigen retrieval was performed by autoclaving at 121°C for 20 min in 10 mM sodium citrate (pH 6), blocked with 5% normal donkey serum in PBS plus Tween 20 (PBST) for 1 hour, then incubated overnight at 4°C in primary antisera diluted in 1% BSA in PBST. Sections were then washed in PBST and incubated with a 1:200 dilution of secondary antybodies for 1 h (Supplemental Table S4). Afterwards, the sections were washed in PBST, counterstained for 20 min with 5 ng/ml DAPI (Thermo Fisher Sci., Waltham, MA), and mounted under coverslips with Prolong Glass (Thermo Fisher).

Earlier version: For immunofluorescent staining frozen sections were used. Organs were collected in tissue freezing medium (Leica Biosystems, Germany). Sections were cut 3 μm thick in the cryostat Leica CM1850 at −20°C. Slides were incubated for 2 min in 4% PFA and washed with PBS three times for 10 min. Sections were permeabilized in 0.03% Triton-X100 (VWR Life Science, Radnor, PA), 0.1% rabbit serum diluted in PBS at 4oC overnight and washed again with PBS. Sections were stained with primary antibodies (Supplementary Table S3) at a dilution of 1:100 for an hour, washed in PBS, and incubated with secondary antibodies (1:100). Slides were stained with 5 ng/ml DAPI (Thermo Fisher Sci., Waltham, MA), mounted with Mowiol® 4-88 (Sigma-Aldrich) and examined on Leica STELLARIS confocal microscope (Leica Microsystems, Germany).

Revised version: For immunofluorescent staining frozen sections organs were collected in tissue freezing medium (Leica Biosystems, Germany). Sections were cut 3 μm thick in the Leica CM1850 cryostat at −20°C. Slides were incubated for 2 min in 4% PFA and washed with PBS three times for 10 min. Sections were permeabilized in 0.03% Triton-X100 (VWR Life Science, Radnor, PA), 0.1% rabbit serum diluted in PBS at 4оC overnight and washed again with PBS. Sections were stained with primary antibodies (Supplemental Table S4) at a dilution of 1:100 for an hour, washed in PBS, and incubated with secondary antibodies (1:100). Slides were stained with 5 ng/ml DAPI, mounted with Mowiol® 4-88 (Sigma-Aldrich) and examined on Leica STELLARIS confocal microscope (Leica Microsystems, Germany).

Added in the revised version: For Oil Red O staining frozen sections (10μm) were used as described above. Staining was carried out according to the manufacturer’s protocol (Sigma, MAK194).

Earlier version: Flow cytometry spermatogenesis assay

Revised version: Flow cytometry

Earlier version: The protocol modified from (Jeyaraj et al., 2003) was used. In testicular cells, this method allows to distinguish between 4n primary spermatocytes that entered meiosis, 2n cells (Leidyg, Sertoli and other somatic cells, spermatogonia stem cells and secondary spermatocytes, that completed first meiotic division) and two 1n groups: round spermatids and elongated spermatids after DNA compaction as well as apoptotic (subG1) cells (Fig. 2A, left) (Jeyaraj et al., 2003; Pereira et al., 2016). The testes of mature mice were placed in phosphate buffered saline (PBS), the tunica albuginea was opened, and the contents were transferred to a solution of 0.5 µg/ml collagenase IV (PanEco, Russia) in PBS and incubated for 15 min at 32°C with shaking at 220 rpm. The seminiferous tubules were washed twice with 1 U DNAse (NE Biolabs, Ipswich, MA) in PBS, transferred to 0.01% trypsin (PanEco), and shaked for 15 min at 32°C, 220 rpm. Trypsin was inactivated with 0.01% bovine serum albumin in PBS; cells were thoroughly resuspended, passed through 70 μm nylon cages (Wuxi NEST Biotechnology, China) and reconstituted in 1 ml PBS. For flow cytometry cells were fixed with 0.75% PFA for 15 minutes at 37°C, then washed in PBS. Cells were lysed in the buffer containing 50 μg/ml propidium iodide (PI), 100 μg/ml RNAse A, 0.1% sodium citrate, 0.3% NP-40 (VWR Life Science) for 30 min at 4°C in the dark followed by flow cytometry analysis on a CytoFlex 26 (Beckman Coulter, Indianapolis, IN) in PE-A and PerCP-A channels. At least 10,000 fluorescent ‘events’ were collected per each sample. Data were analyzed using CytExpert software (Beckman Coulter).

Revised version: For the isolation of the cells from the seminiferous tubules we used the modified protocol from (Jeyaraj, Grossman, and Petrusz 2003). In testicular cells, this method allows to distinguish between 4n primary spermatocytes that entered meiosis, 2n cells (Leidyg, Sertoli and other somatic cells, spermatogonia stem cells and secondary spermatocytes, that completed first meiotic division) and two 1n groups: round spermatids and elongated spermatids after DNA compaction as well as apoptotic (subG1) cells (Fig. 3A, left) (Pereira et al. 2016; Jeyaraj, Grossman, and Petrusz 2003). The testes of mature mice were placed in phosphate buffered saline (PBS), the tunica albuginea was opened, and the contents were transferred to a solution of 0.5 µg/ml collagenase IV (PanEco, Russia) in PBS and incubated for 15 min at 32°C with shaking at 220 rpm. The seminiferous tubules were washed twice with 1 U DNAse (NE Biolabs, Ipswich, MA) in PBS, transferred to 0.01% trypsin (PanEco), and shaken for 15 min at 32°C, 220 rpm. Trypsin was inactivated with 0.01% bovine serum albumin in PBS; cells were thoroughly resuspended, passed through 70 μm nylon cages (Wuxi NEST Biotechnology, China) and reconstituted in 1 ml PBS. For flow cytometry cells were fixed with 0.75% PFA for 15 minutes at 37°C, then washed in PBS. Cells were lysed in the buffer containing 50 μg/ml propidium iodide (PI), 100 μg/ml RNAse A, 0.1% sodium citrate, 0.3% NP-40 (VWR Life Science) for 30 min at 4°C in the dark followed by flow cytometry analysis on a CytoFlex 26 (Beckman Coulter, Indianapolis, IN) in PE-A and PerCP-A channels. At least 10,000 fluorescents ‘events’ were collected per sample.

Added in the revised version: For antibody staining for flow cytometry analysis cells were fixed with 90% methanol for 30 minutes and washed three times with PBS. To reduce potential nonspecific antibody staining Mouse Fc Block (CD16/CD32) (Cell Signaling, #88280) was used as described by the manufacturer. After that, anti-CDK8 (1:100), anti-cyclin C (1:100) and anti-c-Kit (1:20) antibody (Supplemental table S4) diluted in 1% BSA in PBS were added. Cells were incubated for 1 hour, washed three times in PBS, incubated with anti-rabbit Alexa-488 conjugated secondary antibody for 1 hour (Supplemental table S4) and again washed with PBS for three times. For ploidy measurement cells stained only for CDK8 or CcnC were counterstained with PI as described above (Fig. 2E,F and 2H,I respectively), whereas CDK8+c-Kit or CcnC+c-Kit stained cells were counterstained with 5 μg/ml Hoechst 33342 (Invitrogen, H1399) (Fig. 2G and 2J) for 30 min at 4°C in the dark. Flow cytometry measurement using CytoFlex 26. At least 100,000 fluorescents ‘events’ were collected per each sample.

Added in the revised version: All the data were analyzed using CytExpert software (Beckman Coulter).

10x Chromium library preparation and sequencing

Earlier version: 2 R26-Cre-ERT2 and 2 Cdk8fl/flCdk19-/-R26-Cre-ERT2 tamoxifen-treated male mice at 7th week after activation were sacrificed and tissue cell suspension was prepared as described above. The phenotype was confirmed by cell cycle analysis of a portion of cell suspension used for the 10X library preparation. Flow cytometry analysis showed complete absence of haploid cells in one DKO animal while the other had a subpopulation of round but not elongated spermatids (Supplementary Fig. S2).

Revised version: 2 R26-Cre-ERT2 and 2 Cdk8fl/flCdk19−/−R26-Cre-ERT2 tamoxifen-treated male mice at 7th week after activation were sacrificed and tissue cell suspension was prepared as described above. The phenotype was confirmed by cell cycle analysis of a portion of cell suspension used for the 10X library preparation. Flow cytometry analysis showed complete absence of haploid cells in one iDKO animal while the other had a subpopulation of round but not elongated spermatids (Supplemental Fig. S2).

Single cell 3’ v3 kit and Chromium controller (10X Genomics) (Zheng et al. 2017) were used to generate GEMs for further processing. RT and other stages of single cell library preparation were performed as per manufacturer instructions, with the exception of AMPure XP Reagent used instead of SPRIselect. These sample libraries were sequenced with Illumina Novaseq6000.

Single cell data processing

Earlier version: Raw fastq files generation, alignment and read filtering were done by 10x Genomics Cell Ranger 6.1.1 with default settings. Read quality was additionally assessed with FastQC. Reads were aligned to the GRCm38 genome, with an average 95% alignment rate. Raw counts, generated by Cell Ranger 6.1.1, were piped to Seurat (Hao et al., 2021), where additional filtering was applied. Cells with aberrant UMI vs. genes relations were filtered with pagoda2 genes vs. molecule filter (https://github.com/kharchenkolab/pagoda2), where minimal counts were set as 1000 and minimal genes as 500. Next, cells with high percentages of mitochondrial genes expressions per cell were excluded, threshold for filtering set as 25%. After that Scrublet doublet (Wolock et al., 2019) score was calculated, cells with scores higher than 0.20 were marked as doublets and excluded. Each sample had more than 5000 high quality cells remaining after filtering. These cells were normalized with the default Seurat method and piped to further filtering and analysis.

Revised version: Raw fastq files generation, alignment and read filtering were done by 10x Genomics Cell Ranger 6.1.1 with default settings. Read quality was additionally assessed with FastQC. Reads were aligned to the GRCm38 genome, with an average 95% alignment rate. Raw counts, generated by Cell Ranger 6.1.1, were piped to Seurat (Hao et al. 2021), where additional filtering was applied. Cells with aberrant UMI vs. genes relations were filtered with pagoda2 genes vs. molecule filter (https://github.com/kharchenkolab/pagoda2), where minimal counts were set as 1000 and minimal genes as 500. Next, cells with high percentages of mitochondrial genes expressions per cell were excluded, threshold for filtering set as 25%. After that Scrublet doublet (Wolock, Lopez, and Klein 2019) score was calculated, cells with scores higher than 0.20 were marked as doublets and excluded. Each sample had more than 5000 high quality cells remaining after filtering. These cells were normalized with the default Seurat method and piped to further filtering and analysis.

Earlier version: To account for ambient gene expression, DecontX (Yang et al., 2020) was used, for which cluster-assigned data is required. For stable cluster assignment, a reference dataset was formed of WT samples, which were integrated via Harmony (Korsunsky et al., 2018) and then clustered with Seurat. Clustering was checked against testis cell type markers (Cao et al., 2021; Hermann et al., 2018) to ensure selection of true clusters that would reiterate across samples. Cluster of cells belonging to the IGB_2 WT sample only, which did not correspond to any cell type and was missing essential housekeeping genes, was discarded as a cluster of low quality cells with no biological relevance. This reference dataset was used for automatic cluster mapping with Seurat for all samples, and received labels were used with DecontX to lessen share of ambient gene expression in data.

Revised version: To account for ambient gene expression, DecontX (Yang et al. 2020) was used, for which cluster-assigned data is required. For stable cluster assignment, a reference dataset was formed of WT samples, which were integrated via Harmony (Korsunsky et al. 2018) and then clustered with Seurat. Clustering was checked against testis cell type markers (Hermann et al. 2018; Cao et al. 2021) to ensure selection of true clusters that would reiterate across samples. Cluster of cells belonging to the IGB_2 WT sample only, which did not correspond to any cell type and was missing essential housekeeping genes, was discarded as a cluster of low-quality cells with no biological relevance. This reference dataset was used for automatic cluster mapping with Seurat for all samples, and received labels were used with DecontX to lessen share of ambient gene expression in data.

Single cell data analysis

Earlier version: Corrected counts for all cells were piped back to Seurat. Then normalized WT and double knockout samples were integrated with Harmony to account for samples’ batch effects. 3480 variable genes were used for PCA and following integration, none of them mitochondrial or ribosomal. Harmony-integrated data was UMAP dimension-reduced and clustered with Seurat. Visualizations were also made with the Seurat package. Cell type was assigned according to their marker genes (Cao et al., 2021; Hermann et al., 2018). Both cells and microenvironment in the testis were detected in wild-type as well as in KO samples.

Revised version: Corrected counts for all cells were piped back to Seurat. Then normalized WT and double knockout samples were integrated with Harmony to account for samples’ batch effects. 3480 variable genes were used for PCA and following integration, none of them mitochondrial or ribosomal. Harmony-integrated data was UMAP dimension-reduced and clustered with Seurat. Visualizations were also made with the Seurat package. Cell type was assigned according to their marker genes (Hermann et al. 2018; Cao et al. 2021). Both cells and microenvironment in the testis were detected in wild-type as well as in KO samples.

Earlier version: Spermatogonia cell cluster was reintegrated and reclustered separately for discovery of spermatogonia subtypes. Integration and clusterization was done as before, with 798 features at the start of the process. Spermatogonia subclusters were assigned types according to literature-curated marker genes (Cao et al., 2021; Hermann et al., 2018). Differential expression tests (negative binomial Seurat implementation) were conducted for each chosen cell cluster with WT compared against double knockout, batch effect between samples accounted for with the linear model. All mitochondria, ribosomal and high ambient genes (by DecontX) were excluded from testing. Genes with absolute logFC < 0.4 and genes expressed in less than 20% cells in class were excluded either. Additionally, after testing, all genes with no or almost no expression in double knockout, high expressions in WT, and high expression in WT spermatids were excluded as their differential expression might be caused by different ambient composition in WT, heavily influenced by larger spermatid cell proportion. For Sertoli cells, a small subcluster, potentially caused by technical effects, was excluded before testing.

Revised version: Spermatogonia cell cluster was reintegrated and reclustered separately for discovery of spermatogonia subtypes. Integration and clusterization was done as before, with 798 features at the start of the process. Spermatogonia subclusters were assigned types according to literature-curated marker genes (Hermann et al. 2018; Cao et al. 2021). Differential expression tests (negative binomial Seurat implementation) were conducted for each chosen cell cluster with WT compared against double knockout, batch effect between samples accounted for with the linear model. All mitochondria, ribosomal and high ambient genes (by DecontX) were excluded from testing. Genes with absolute logFC < 0.4 and genes expressed in less than 20% cells in class were excluded either. Additionally, after testing, all genes with no or almost no expression in double knockout, high expressions in WT, and high expression in WT spermatids were excluded as their differential expression might be caused by different ambient composition in WT, heavily influenced by larger spermatid cell proportion. For Sertoli cells, a small subcluster, potentially caused by technical effects, was excluded before testing.

Assignment of groups and cell count

Earlier version: The clustering analysis was performed according to markers described in previous studies (reviewed in (Suzuki, 2023)). This analysis allowed us to identify nine populations with respective genes: Sertoli cells (Wt1, Sox9, NR5a1, Clu), Leydig cells (Hsd3b1, Cyp17a1), macrophages (Cd74, C1qa, Cd45, Ccl5), T-cells, fibroblasts (Pdfrgfra, Cd34, Dcn, Gcn), and the differentiating sperm lineage divided into 5 clusters: undifferentiated/early spermatogonia (Crabp1, Tcea3, Kit), meiotic entry (Stra8, Rhox13) leptotene/zygotene (Sycp3, Dmc1), two clusters for pachytene (Piwil1, Id4) and pachytene/diplotene (Acr, Pgk2) (Supplementary Fig. S5).

Revised version: The clustering analysis was performed according to markers described in previous studies (reviewed in (Suzuki 2023)). This analysis allowed us to identify nine populations with respective genes: Sertoli cells (Wt1, Sox9, NR5a1, Clu), Leydig cells (Hsd3b1, Cyp17a1), macrophages (Cd74, C1qa, Cd45, Ccl5), T-cells, fibroblasts (Pdfrgfra, Cd34, Dcn, Gcn), and the differentiating sperm lineage divided into 5 clusters: undifferentiated/early spermatogonia (Crabp1, Tcea3, Kit), meiotic entry (Stra8, Rhox13) leptotene/zygotene (Sycp3, Dmc1), two clusters for pachytene (Piwil1, Id4) and pachytene/diplotene (Acr, Pgk2) (Supplementary Fig. S7).

Hormone measurement

To measure the concentration of hormones in the serum, blood was taken from the left ventricle of the heart, transferred to sterile test tubes and left at 4°C overnight. Measurement of luteinizing hormone was performed using ELISA (CEA441Mu, CloudClone, China). The testosterone concentration was measured by the LC-MS method as described previously (Povaliaeva et al. 2020). 17α-Hydroxyprogesterone-d8 were used as standard, proteins were precipitated by adding ZnSO4 and MeOH. The measurement was performed using an on-line extraction method with Agilent Bond Elut C18 cartridges as trap column and Waters Aquity UPLC BEH C18 column, Agilent 1290 Infinity II LC and and AB Sciex Triple Quad 5500 mass-spectrometer.

Added in the revised version: Generation of antibodies against CDK19

Added in the revised version: Coding sequences for aa 377–473 of CDK19 (Q8BWD8) were cloned into pGEX5.1 expression vectors. Recombinant GST-tagged CDK19 epitope proteins were purified with Glutathione Sepharose and used for immunization of rabbits to generate target-specific polyclonal antibodies. Animals were immunized every two weeks with an intradermal injection of a Complete Freund’s Adjuvant (Sigma, F5881) and 500 μg of recombinant CDK19. After 6-8 months of immunization, blood was collected and antibodies were purified with BrCN-activated Sepharose (Pharmacia).

Immunoblotting

Earlier version: Protein lysis and western blotting were performed as described in (Ilchuk et al., 2022). Briefly, cells were lysed in RIPA buffer supplemented with protein inhibitor cocktail (Sigma-Aldrich, St. Louis, MO). Total protein concentration was quantified by the Bradford method. Absorbance at 560 nm was measured with a CLARIOstar Plate Reader (BMG Labtech, Germany). Proteins were separated by SDS-PAGE and transferred onto 0.2 μm nitrocellulose membrane (Bio-Rad, Hercules, CA). After blocking with 5% skimmed milk, membranes were treated with primary antibodies (Supplementary table S4) and incubated at 4oC overnight. Membranes were washed with Tris-borate saline with Tween 20 (TBS-T) and incubated for 1 h at room temperature with secondary antibodies (Supplementary table S4). Membranes were visualized with the Clarity Western ECL Substrate (Bio-Rad) using a iBright FL1500 Imaging System (Invitrogen, Waltham, MA).

Revised version: Protein lysis and western blotting were performed as described in (Ilchuk et al. 2022). Briefly, cells were lysed in RIPA buffer supplemented with protein inhibitor cocktail (Sigma-Aldrich, St. Louis, MO). Total protein concentration was quantified by the Bradford method. Absorbance at 560 nm was measured with a CLARIOstar Plate Reader (BMG Labtech, Germany). Proteins were separated by SDS-PAGE and transferred onto 0.2 μm nitrocellulose membrane (Bio-Rad, Hercules, CA). After blocking with 5% skimmed milk, membranes were treated with primary antibodies (Supplemental table S4) and incubated at 4°C overnight. Membranes were washed with Tris-borate saline with Tween 20 (TBS-T) and incubated for 1 h at room temperature with secondary antibodies (Supplemental table S4). Membranes were visualized with the Clarity Western ECL Substrate (Bio-Rad) using a iBright FL1500 Imaging System (Invitrogen, Waltham, MA).

Leydig cell culture

Earlier version: Seminiferous tubules were isolated as described above for flow cytometry but after collagenase IV (1 ug/ml) digestion the supernatant was passed through 70 μm nylon cages and centrifuged at 1000 rpm. The cells were resuspended in the medium for Leydig cells (low glucose DMEM (PanEco) with 10% fetal bovine serum (FBS, HyClone, GE Healthcare Life Sciences, Chicago, IL), 1% penicillin/streptomycin (PanEco), 0.3 mg/ml L-glutamine (PanEco)) and seeded on Costar® 6-well plates (Corning, New York, NY) pre-coated by 0.01% poly-L-lysyne (Sigma-Aldrich, P0899), incubated at 37°C in 5% CO2 for 1 h and carefully washed with the medium. After 24 h cells were treated with 0.56% potassium chloride hypotonic solution for 5 min to remove unattached cells. On the 7th day 1 µM Senexin B (SenB, Senex Biotechnology, Columbia, SC) or 1 µM 4-hydroxytamoxifen (4-OHT, Sigma-Aldrich) was added to the cells. After 7 days of incubation, the cells were lysed with ExtractRNA (Evrogen, Russia).

Revised version: Seminiferous tubules were isolated as described above for flow cytometry but after collagenase IV (1 µg/ml) digestion the supernatant was passed through 70 μm nylon cages and centrifuged at 1000 rpm. The cells were resuspended in the medium for Leydig cells (low glucose DMEM (PanEco) with 10% fetal bovine serum (FBS, HyClone, GE Healthcare Life Sciences, Chicago, IL), 1% penicillin/streptomycin (PanEco), 0.3 mg/ml L-glutamine (PanEco)) and seeded on Costar® 6-well plates (Corning, New York, NY) pre-coated by 0.01% poly-L-lysyne (Sigma-Aldrich, P0899), incubated at 37°C in 5% CO2 for 1 h and carefully washed with the medium. After 24 h cells were treated with 0.56% potassium chloride hypotonic solution for 5 min to remove unattached cells. On the 7th day 1 µM Senexin B (SenB, Senex Biotechnology, Columbia, SC) or 1 µM 4-hydroxytamoxifen (4-OHT, Sigma-Aldrich) was added to the cells. After 7 days of incubation, the cells were lysed with ExtractRNA (Evrogen, Russia).

qRT PCR

Earlier version: Total RNA from cell suspension was extracted with ExtractRNA (Evrogen) and quantified using NanoDrop. Equal amounts of RNA were reverse-transcribed to generate cDNA using Superscript II reverse transcription Supermix for qRT-PCR (Invitrogen). qRT-PCR was then performed with SYBR PCR Mix (Evrogen) using the Thermal Cycler system (Bio-Rad). The data thus obtained were analyzed following the comparative (ΔΔCt) method. Actb was used as a housekeeping gene. Sequences of primers used are listed in Supplementary table S4.

Revised version: Total RNA from cell suspension was extracted with ExtractRNA (Evrogen) and quantified using NanoDrop. Equal amounts of RNA were reverse-transcribed to generate cDNA using Superscript II reverse transcription Supermix for qRT-PCR (Invitrogen). qRT-PCR was then performed with SYBR PCR Mix (Evrogen) using the Thermal Cycler system (Bio-Rad). The data thus obtained were analyzed following the comparative (ΔΔCt) method. Actb was used as a housekeeping gene. Sequences of primers used are listed in Supplemental table S4.

Statistics

Earlier version: The normality of data was tested using the Shapiro–Wilk test, all datasets met the condition for normality (p > 0.05). One-way or two-way analysis of variance (ANOVA) followed by Holm-Sidak’s post hoc test for multiple comparisons was used (GraphPad Prism 8; GraphPad Software, San Diego, CA). P value <0.05 was taken as evidence of statistical significance.

Revised version: The normality of data was tested using the Shapiro–Wilk test, all datasets met the condition for normality (pL>L0.05). One-way or two-way analysis of variance (ANOVA) followed by Holm-Sidak’s post hoc test for multiple comparisons was used (GraphPad Prism 8; GraphPad Software, San Diego, CA). P value <0.05 was taken as evidence of statistical significance. 114 mice were used for experiments described here: 34 wild type and Rosa-Cre/ERT2 control mice, 16 CDK8 iKO mice, 15 CDK19 KO mice and 46 iDKO mice. All experiments had at least two biological repeats.

Removed in the revised version: Declaration of interests

Removed in the revised version: I.B.R is Founder and President and M.C. is consultant of Senex Biotechnology, Inc. Other authors declare no conflict of interest.

Removed in the revised version: Supporting information